```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Homework 02: Arrays and Tables

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided modules. Each time you start your server, you will need to execute this cell again to load the modules.

```python
# Don't change this cell; just run it. 
import numpy as np
from datascience import *
```

**Helpful Resource:**
- [Python Reference](http://data8.org/sp21/python-reference.html): Cheat sheet of helpful array & table methods used in this course.

**Recommended Reading:**
- [What is Data Science](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)

- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 

- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!



**Notes:** 

- Start early so that you can come to office hours if you're stuck. Check the course website for the office hours schedule. Late work will not be accepted as per the course expectations. Late work will not be accepted without advanced permission as per the course expectations document, so plan accordingly. You should start early so that you have time to get help if you're stuck. Office hours are held throughout the week, so be sure to use them if needed.

- This homework has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. Additional tests will be run once your homework is submitted to Gradescope**. Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and not super far off the expected answer. The tests that are run after submission will evaluate accuracy more carefully. Do not assume that just because all your tests pass means that your answers are correct!

- Directly sharing answers is **not** okay, but discussing problems with the course staff (teacher and TAs) or with other students is encouraged. Refer to the course expectations document to learn more about how to learn and work cooperatively.

## 1. Creating Arrays

**Question 1.** Make an **array** called `weird_numbers` containing the following numbers (in the given order):

1. $-2$
2. $\sin(12)$
3. $3$
4. $5^{\cos(12)}$

**Hint:** `sin` and `cos` are functions in the `math` module. Importing modules is covered in a previous lab! You should import the module in your solution cell below.

**Note:** Python lists are different/behave differently than numpy arrays. In this course, we use `numpy` arrays, so please make an **array**, not a python list.


```python
# Our solution involved one extra line of code before creating
# weird_numbers.
...
weird_numbers = ...
weird_numbers
```

```python
grader.check("q1")
```

**Question 2.** Make an array called `book_title_words` containing the following three strings: "Eats", "Shoots", and "and Leaves".


```python
book_title_words = ...
book_title_words
```

```python
grader.check("q2")
```

### `join`
Strings have a method called `join`.  `join` takes one argument, an array of strings.  It returns a single string.  Specifically, the value of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, **except** `a_string` is inserted in between each string in the array.

**Question 3.** Use the array `book_title_words` and the method `join` to make two strings:

1. "Eats, Shoots, and Leaves" (call this one `with_commas`)

2. "Eats Shoots and Leaves" (call this one `without_commas`)

**Hint:** If you're not sure what `join` does, first try just calling, for example, `"foo".join(book_title_words)` .


```python
with_commas = ...
without_commas = ...

print('with commas: ', with_commas)
print('without commas: ', without_commas)
```

```python
grader.check("q3")
```

## 2. Indexing Arrays

These exercises give you practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *index*, so the first element is the element at index 0.

**Note:** If you have previous coding experience, you may be familiar with bracket notation. **DO NOT** use bracket notation when indexing (i.e. `arr[0]`), as this can yield different data type outputs than what we will be expecting. This can cause you to fail an autograder test.**

Be sure to refer to the [Python Reference](http://data8.org/fa20/python-reference.html) on the website if you feel stuck!

**Question 4.** The cell below creates an array of some numbers.  Set `third_element` to the third element of `some_numbers`.


```python
some_numbers = make_array(-1, -3, -6, -10, -15)

third_element = ...
third_element
```

```python
grader.check("q4")
```

**Question 5.** The next cell creates a table that displays some information about the elements of `some_numbers` and their order.  Run the cell to see the partially-completed table, then fill in the missing information (the cells that say "Ellipsis") by assigning `blank_a`, `blank_b`, `blank_c`, and `blank_d` to the correct elements in the table.

**Hint:** Replace the `...` with strings or numbers.


```python
blank_a = ...
blank_b = ...
blank_c = ...
blank_d = ...

# Don't change anything below this line!
elements_of_some_numbers = Table().with_columns(
    "English name for position", make_array("first", "second", blank_a, blank_b, "fifth"),
    "Index",                     make_array(blank_c, 1, 2, blank_d, 4),
    "Element",                   some_numbers)
elements_of_some_numbers
```

```python
grader.check("q5")
```

**Question 6.** You'll sometimes want to find the *last* element of an array.  Suppose an array has 142 elements.  What is the index of its last element?


```python
index_of_last_element = ...
```

```python
grader.check("q6")
```

More often, you don't know the number of elements in an array, its *length*.  (For example, it might be a large dataset you found on the Internet.)  The function `len` takes a single argument, an array, and returns the `len`gth of that array (an integer).

**Question 7.** The cell below loads an array called `president_birth_years`.  Calling `.column(...)` on a table returns an array of the column specified, in this case the `Birth Year` column of the `president_births` table. The last element in that array is the most recent birth year of any deceased president in this data set. Assign that year to `most_recent_birth_year`. Your solution should used both the `len(...)` and `.item(...)` functions, and you should not have to know the number of rows in `president_birth_years` to answer this question. Your code should be able to work no matter how big or small `president_birth_years` might be.


```python
president_birth_years = Table.read_table("data/president_births.csv").column('Birth Year')

most_recent_birth_year = ...
most_recent_birth_year
```

```python
grader.check("q7")
```

**Question 8.** Finally, assign `sum_of_birth_years` to the sum of the first, sixteenth, and last birth year in `president_birth_years`. Your solution should not include the actual years in the calculation, but rather references to the array where those years are stored.


```python
sum_of_birth_years = ...
sum_of_birth_years
```

```python
grader.check("q8")
```

## 3. Basic Array Arithmetic

**Question 9.** Multiply the numbers 42, $-4224$, $424224242$, and 250 by 157. Assign each variable below such that `first_product` is assigned to the result of $42 \cdot 157$, `second_product` is assigned to the result of $-4224 \cdot 157$, and so on. 

For this question, **don't** use arrays.


```python
first_product = ...
second_product = ...
third_product = ...
fourth_product = ...
print(first_product, second_product, third_product, fourth_product)
```

```python
grader.check("q9")
```

**Question 10.** Now, do the same calculation, but using an array called `numbers` that contains the 4 original numbers (42, $-4224$, $424224242$, and 250) and only a single multiplication (`*`) operator.  Store the 4 results in an array named `products`.


```python
numbers = ...
products = ...
products
```

```python
grader.check("q10")
```

**Question 11.** Oops, we made a typo!  Instead of 157, we wanted to multiply each number by 1577.  Compute the correct products in the cell below using array arithmetic.  Notice that your job is really easy if you previously defined an array containing the 4 numbers.


```python
correct_products = ...
correct_products
```

```python
grader.check("q11")
```

**Question 12.** We've loaded an array of temperatures in the next cell.  Each number is the highest temperature observed on a day at a climate observation station, mostly from the US.  Since they're from the US government agency [NOAA](https://noaa.gov), all the temperatures are in Fahrenheit.  Convert them all to Celsius by first subtracting 32 from them, then multiplying the results by $\frac{5}{9}$. Make sure to **round** the final result after converting to Celsius to the nearest integer using the `np.round` function. `np.round` works a lot like the builtin `round` function in python, but it can operate on arrays.


```python
max_temperatures = Table.read_table("data/temperatures.csv").column("Daily Max Temperature")

celsius_max_temperatures = ...
celsius_max_temperatures
```

```python
grader.check("q12")
```

**Question 13.** The cell below loads all the *lowest* temperatures from each day (in Fahrenheit).  Compute the size of the daily temperature range for each day.  That is, compute the difference between each daily maximum temperature and the corresponding daily minimum temperature.  **Pay attention to the units, give your answer in Celsius.** Make sure **NOT** to round your answer for this question! Note: Remember that in **Question 12**, `celsius_max_temperatures` was rounded, so you probably don't want to use that in this question.


```python
min_temperatures = Table.read_table("data/temperatures.csv").column("Daily Min Temperature")

celsius_temperature_ranges = ...
celsius_temperature_ranges
```

```python
grader.check("q13")
```

## 4. World Population


The tests from this point on will **not** necessarily tell you whether or not your answers are correct.

The cell below loads a table of estimates of the world population for different years, starting in 1950. The estimates come from the [US Census Bureau website](https://www.census.gov/en.html).

```python
world = Table.read_table("data/world_population.csv").select('Year', 'Population')
world.show(4)
```

The name `population` is assigned to an array of population estimates.

```python
population = world.column('Population')
population
```

In this question, you will apply some built-in Numpy functions to this array. Numpy is a module that is often used in Data Science!

<img src="images/array_diff.png" style="width: 600px;"/>

The difference function `np.diff` subtracts each element in an array from the element after it within the array. As a result, the length of the array `np.diff` returns will always be one less than the length of the input array.

<img src="images/array_cumsum.png" style="width: 700px;"/>

The cumulative sum function `np.cumsum` outputs an array of partial sums. For example, the third element in the output array corresponds to the sum of the first, second, and third elements.

**Note:** This homework aims to get you comfortable with array arithmetic using methods like `np.diff` and `np.cumsum`, but they are not heavily used later in the course.

**Question 14.** Very often in data science, we are interested understanding how values change with time. Use `np.diff` and `np.max` (or just `max`) to calculate the largest annual change in population between any two consecutive years.


```python
largest_population_change = ...
largest_population_change
```

```python
grader.check("q14")
```

**Question 15.** What do the values in the resulting array represent? Choose one of the following options, and assign the corresponding number to the name `cumulative_sum_answer`.

**Tip:** Look at the population array and compute the actual population changes for the first couple years of each option.

```python
np.cumsum(np.diff(population))
```

1. The total population change between consecutive years, starting at 1951.

1. The total population change between 1950 and each later year, starting at 1951.

1. The total population change between 1950 and each later year, starting inclusively at 1950.


```python
# Assign cumulative_sum_answer to 1, 2, or 3
cumulative_sum_answer = ...
```

```python
grader.check("q15")
```

## 5. Old Faithful


Old Faithful is a geyser in Yellowstone that erupts every 44 to 125 minutes (according to [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). People are [often told that the geyser erupts every hour](http://yellowstone.net/geysers/old-faithful/), but in fact the waiting time between eruptions is more variable. Let's take a look.

**Question 16.** The first line below assigns `waiting_times` to an array of 272 consecutive waiting times between eruptions, taken from a classic 1938 dataset. Assign the names `shortest`, `longest`, and `average` so that the `print` statement is correct.


```python
waiting_times = Table.read_table('data/old_faithful.csv').column('waiting')

shortest = ...
longest = ...
average = ...

print("Old Faithful erupts every", shortest, "to", longest, "minutes and every", average, "minutes on average.")
```

```python
grader.check("q16")
```

**Question 17.** Assign `biggest_decrease` to the biggest decrease in waiting time between two consecutive eruptions. For example, the third eruption occurred after 74 minutes and the fourth after 62 minutes, so the decrease in waiting time was $62 - 74 = -12$ minutes, for a decrease of 12 minutes.

**Hints:**

- You'll need an array arithmetic function [mentioned in the textbook](https://inferentialthinking.com/chapters/05/1/Arrays.html#functions-on-arrays). You have also seen this function earlier in the homework!

- We want to return the absolute value of the biggest decrease.


```python
biggest_decrease = ...
biggest_decrease
```

```python
grader.check("q17")
```

**Question 18.** If you expected Old Faithful to erupt every hour, you would expect to wait a total of `60 * k` minutes to see `k` eruptions. Use the `np.arange` function to create a range of values where each element represents the total expected time (in minutes) one would spend waiting for  waiting for `k` erruptions to occur under the assumption erumptions occur once an hour. Set `expected_wait` to this range of values. You can read about `np.arange` as needed in this [textbook section](https://inferentialthinking.com/chapters/05/2/Ranges.html).

Then, set `difference_from_expected` to an array with 272 elements, where the element at index `i` is the absolute difference between the actual and expected total amount of waiting time to see the first `i+1` eruptions. 

**Hint 1**: You'll need to compare a cumulative sum to a range. 

For example, since the first three waiting times are 79, 54, and 74, the total waiting time for 3 eruptions is 79 + 54 + 74 = 207. The expected waiting time for 3 eruptions is 60 * 3 = 180. Therefore, `difference_from_expected.item(2)` should be $|207 - 180| = 27$. 

**Hint 2**: When using the absolute value function, use `np.abs` instead of `abs`, since `numpy` is already loaded in this notebook, and that's what the autograder is going to use.


```python
expected_wait = ...
difference_from_expected = ...
difference_from_expected
```

```python
grader.check("q18")
```

**Question 20.** Let’s imagine your guess for the next wait time was always just the length of the previous waiting time. If you always guessed the previous waiting time, how big would your error in guessing the waiting times be, on average?

For example, the first three waiting times are 79, 54, and 74. After the first eruption you guess the next eruption will happen after 79 minutes, but it occurs after 54 resulting in an error of $\lvert 79 - 54 \rvert$. For the next eruption you now guess it will take 54 minutes, but it takes 74, so the error is $\lvert 54 - 74 \rvert$. Therefore, the average difference between your guess and the actual time for just the second and third eruption would be $$\frac{|79-54|+ |54-74|}{2} = 22.5$$.


```python
average_error = ...
average_error
```

```python
grader.check("q20")
```

## 6. Tables


**Question 21.** Suppose you have 4 apples, 3 oranges, and 3 pineapples.  (Perhaps you're using Python to solve a high school Algebra problem.)  Create a table that contains this information.  It should have two columns: `fruit name` and `count`.  Assign the new table to the variable `fruits`.

**Note:** Use lower-case and singular words for the name of each fruit, like `"apple"`.

Your table should look like:

| fruit name | count |
|------------|-------|
| apple      | 4     |
| orange     | 3     |
| pineapple  | 3     |


```python
# Our solution uses 1 statement split over 3 lines. 
# You can write yours on 1 line if you wish.
fruits = ...
    ...
    ...
fruits
```

```python
grader.check("q21")
```

**Question 22.** The file `inventory.csv` contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit. Load it as a table named `inventory` using the `Table.read_table()` function. `Table.read_table(...)` takes one argument (data file name in string format) and returns a table.


```python
inventory = ...
inventory
```

```python
grader.check("q22")
```

**Question 23.** Does each box at the fruit stand contain a different fruit? Set `all_different` to `True` if each box contains a different fruit or to `False` if multiple boxes contain the same fruit.

**Hint:** You don't have to write code to calculate the True/False value for `all_different`. Just look at the `inventory` table and assign `all_different` to either `True` or `False` according to what you can see from the table in answering the question.


```python
all_different = ...
all_different
```

```python
grader.check("q23")
```

**Question 24.** The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called "price per fruit (\$)" that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  

Load these data into a table called `sales`.


```python
sales = ...
sales
```

```python
grader.check("q24")
```

**Question 25.** How many fruits did the store sell in total on that day?


```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("q25")
```

**Question 26.** What was the store's total revenue (the total price of all fruits sold) on that day?

**Hint:** If you're stuck, think first about how you would compute the total revenue from just the grape sales.


```python
total_revenue = ...
total_revenue
```

```python
grader.check("q26")
```

**Question 27.** Make a new table called `remaining_inventory`.  It should have the same rows and columns as `inventory`, except that the amount of fruit sold from each box should be subtracted from that box's count, so that the "count" is the amount of fruit remaining after Saturday. Remember, the rows in `inventory` and `sales` are already in the same order.


```python
remaining_inventory = ...
    ...
    ...
    ...
remaining_inventory
```

```python
grader.check("q27")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> isinstance(weird_numbers, np.ndarray)
True
```
```
>>> len(weird_numbers)
4
```
```
>>> np.allclose(weird_numbers, np.array([-2.0, -0.53657292, 3.0, 3.88891638]), rtol=0.001, atol=0.001)
True
```


Tests for q10 (points: None)
```
>>> np.allclose(products, np.array([6594, -663168, 66603205994, 39250]))
True
```


Tests for q11 (points: None)
```
>>> np.allclose(correct_products, np.array([66234, -6661248, 669001629634, 394250]))
True
```


Tests for q12 (points: None)
```
>>> sum(celsius_max_temperatures) != 356705.0
True
```
```
>>> sum(celsius_max_temperatures)
1280677.0
```
```
>>> len(celsius_max_temperatures)
65000
```
```
>>> celsius_max_temperatures.item(2003)
20.0
```


Tests for q13 (points: None)
```
>>> np.round(sum(celsius_temperature_ranges), 0)
768487.0
```
```
>>> len(celsius_temperature_ranges)
65000
```
```
>>> celsius_temperature_ranges.item(1)
10.0
```


Tests for q14 (points: None)
```
>>> largest_population_change
87515824
```


Tests for q15 (points: None)
```
>>> isinstance(cumulative_sum_answer, int)
True
```
```
>>> cumulative_sum_answer in [1, 2, 3]
True
```


Tests for q16 (points: None)
```
>>> shortest <= average <= longest
True
```


Tests for q17 (points: None)
```
>>> biggest_decrease != 47
True
```
```
>>> 30 <= biggest_decrease < 47
True
```


Tests for q18 (points: None)
```
>>> difference_from_expected.size
272
```
```
>>> difference_from_expected.item(271) == np.abs(60 * 272 - sum(waiting_times))
True
```


Tests for q2 (points: None)
```
>>> not any([',' in text for text in book_title_words])
True
```
```
>>> 'and ' in book_title_words.item(2)
True
```
```
>>> len(book_title_words)
3
```
```
>>> book_title_words
array(['Eats', 'Shoots', 'and Leaves'],
      dtype='<U10')
```


Tests for q20 (points: None)
```
>>> 15 <= average_error <= 25
True
```


Tests for q21 (points: None)
```
>>> fruits.sort(0)
fruit name | count
apple      | 4
orange     | 3
pineapple  | 3
```


Tests for q22 (points: None)
```
>>> inventory.sort(0).column(0).item(0)
25274
```


Tests for q23 (points: None)
```
>>> all_different in {True, False}
True
```


Tests for q24 (points: None)
```
>>> sales.sort(0)
box ID | fruit name | count sold | price per fruit ($)
25274  | apple      | 0          | 0.8
26187  | strawberry | 25         | 0.15
43566  | peach      | 17         | 0.8
48800  | orange     | 35         | 0.6
52357  | strawberry | 102        | 0.25
53686  | kiwi       | 3          | 0.5
57181  | strawberry | 101        | 0.2
57930  | grape      | 355        | 0.06
```


Tests for q25 (points: None)
```
>>> total_fruits_sold > 10
True
```
```
>>> isinstance(total_fruits_sold, np.int64)
True
```


Tests for q26 (points: None)
```
>>> 50 <= total_revenue <= 150
True
```


Tests for q27 (points: None)
```
>>> remaining_inventory.num_columns
3
```
```
>>> remaining_inventory.column('count').item(0) != 45
True
```
```
>>> remaining_inventory.where('fruit name', 'grape')
box ID | fruit name | count
57930  | grape      | 162
```


Tests for q3 (points: None)
```
>>> ',' in with_commas
True
```
```
>>> ',' not in without_commas
True
```
```
>>> len(with_commas)
24
```
```
>>> len(without_commas)
22
```
```
>>> with_commas == 'Eats, Shoots, and Leaves'
True
```
```
>>> without_commas == 'Eats Shoots and Leaves'
True
```


Tests for q4 (points: None)
```
>>> type(third_element) == int or type(third_element) == numpy.int64
True
```
```
>>> third_element != -10
True
```
```
>>> third_element
-6
```


Tests for q5 (points: None)
```
>>> elements_of_some_numbers.column(0).item(2)
'third'
```
```
>>> elements_of_some_numbers.column(0).item(3)
'fourth'
```
```
>>> elements_of_some_numbers.column(1).item(0)
0
```
```
>>> elements_of_some_numbers.column(1).item(3)
3
```


Tests for q6 (points: None)
```
>>> index_of_last_element
141
```


Tests for q7 (points: None)
```
>>> most_recent_birth_year
1924
```


Tests for q8 (points: None)
```
>>> sum_of_birth_years
5464
```


Tests for q9 (points: None)
```
>>> first_product
6594
```
```
>>> second_product
-663168
```
```
>>> third_product
66603205994
```
```
>>> fourth_product
39250
```
