```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

# Run the cell below

To run a code cell (i.e.; execute the python code inside a Jupyter notebook) you can click the play button on the ribbon underneath the name of the notebook. Before you begin click the "Run cell" button at the top that looks like ▶| or hold down `Shift` + `Return`.

# Homework 07: Confidence Intervals

**Reading**: 
* [Estimation](https://www.inferentialthinking.com/chapters/13/estimation.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Notes:** 

- Start early so that you can come to office hours if you're stuck. Check the course website for the office hours schedule. Late work will not be accepted as per the course expectations. Late work will not be accepted without advanced permission as per the course expectations document, so plan accordingly. You should start early so that you have time to get help if you're stuck. Office hours are held throughout the week, so be sure to use them if needed.

- This homework has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. Additional tests will be run once your homework is submitted to Gradescope**. Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and not super far off the expected answer. The tests that are run after submission will evaluate accuracy more carefully. Do not assume that just because all your tests pass means that your answers are correct!

- Directly sharing answers is **not** okay, but discussing problems with the course staff (teacher and TAs) or with other students is encouraged. Refer to the course expectations document to learn more about how to learn and work cooperatively.

```python
# Don't change this cell; just run it. 
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 1. Thai Restaurants


Max and Aesha are trying see what the best Thai restaurant in Durham, NC is. They survey 1500 Durham residents selected uniformly at random, and ask each of them what Thai restaurant is the best. 

**Note:** This data is fabricated for the purposes of this homework, but these restaurants are real if you wanted to try them out! 

The choices of Thai restaurant are Bua Thai, Thai Cafe, Pad Thai, and Thai Spoon. After compiling the results, Max and Aesha release the following percentages from their sample:

|Thai Restaurant  | Percentage|
|:------------:|:------------:|
|Bua Thai | 8% |
|Thai Cafe | 52% |
|Pad Thai | 25% |
|Thai Spoon | 15% |

These percentages represent a uniform random sample of the population of Durham residents. We will attempt to estimate the corresponding *parameters*, or the percentage of the votes that each restaurant will receive from the entire population (the entire population is all Durham residents). We will use confidence intervals to compute a range of values that reflects the uncertainty of our estimates.

The table `votes` contains the results of the survey.

```python
# Just run this cell
votes = Table.read_table('data/votes.csv')
votes
```

**Question 1.** Complete the function `one_resampled_percentage` below. It should return Thai Cafe's **percentage** (not proportion) of votes after simulating one bootstrap sample from `tbl`.

**Note:** `tbl` will always be in the same format as `votes`.

```python
def one_resampled_percentage(tbl):
    ...

one_resampled_percentage(votes)
```

```python
grader.check("q1")
```

**Question 2.** Complete the `percentages_in_resamples` function such that it returns an array of 2500 bootstrapped estimates of the percentage of voters who will vote for Thai Cafe. You should use the `one_resampled_percentage` function you wrote above.


```python
def percentages_in_resamples():
    percentage_thai_cafe = make_array()
    ...
percentages_in_resamples()
```

```python
grader.check("q2")
```

In the following cell, we run the function you just defined, `percentages_in_resamples`, and create a histogram of the calculated statistic for the 2,500 bootstrap estimates of the percentage of voters who voted for Thai Cafe. Based on what the original Thai restaurant percentages were, does the graph seem reasonable? Talk to a friend or ask your teacher if you are unsure! You don't need to provide a written response but make sure you've thought about it before moving on.

```python
resampled_percentages = percentages_in_resamples()
Table().with_column('Estimated Percentage', resampled_percentages).hist("Estimated Percentage")
```

**Question 3.** Using the array `resampled_percentages` that was created in the previous code cell, find the values at the two edges of the middle 95% of the bootstrapped percentage estimates. (Compute the lower and upper ends of the interval, named `thai_cafe_lower_bound` and `thai_cafe_upper_bound`, respectively.)


```python
thai_cafe_lower_bound = ...
thai_cafe_upper_bound = ...

# The following code wil print the results of your calculations.
# Don't change it, just run it!
print("Bootstrapped 95% confidence interval for the percentage of Thai Cafe voters in the population: [{:f}, {:f}]".format(thai_cafe_lower_bound, thai_cafe_upper_bound))
```

```python
grader.check("q3")
```

**Question 4.** The survey results seem to indicate that Thai Cafe is beating all the other Thai restaurants combined among voters. We would like to use confidence intervals to determine a range of likely values for Thai Cafe's true lead over all the other restaurants combined. The calculation for Thai Cafe's lead over Bua Thai, Pad Thai, and Thai Spoon combined is:

$$\text{Thai Cafe's % of the vote} - \text{(Bua Thai's % of the vote + Pad Thai's % of the vote + Thai Spoons's % of the vote)}$$

Define the function `one_resampled_difference` that returns **exactly one value** of Thai Cafe's percentage lead over Bua Thai, Pad Thai, and Thai Spoon combined from one bootstrap sample of `tbl`.


```python
def one_resampled_difference(tbl):
    bootstrap = ...
    tc_percentage = ...
    bt_percentage = ...
    pt_percentage = ...
    ts_percentage = ...
    ...

one_resampled_difference(votes)
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

**Question 5.** Write a function called `leads_in_resamples` that finds 2,500 bootstrapped estimates (the result of calling `one_resampled_difference`) of Thai Cafe's lead over Bua Thai, Pad Thai, and Thai Spoon combined. Plot a histogram of the resulting samples. 

**Note:** Thai Cafe's lead can be negative.


```python
def leads_in_resamples():
    ...

sampled_leads = leads_in_resamples()
Table().with_column('Estimated Lead', sampled_leads).hist("Estimated Lead")
```

<!-- END QUESTION -->

**Question 6.** Use the simulated data from Question 1.5 to compute an approximate 95% confidence interval for Thai Cafe's true lead over Bua Thai, Pad Thai, and Thai Spoon combined.


```python
diff_lower_bound = ...
diff_upper_bound = ...

# The following code wil print the results of your calculations.
# Don't change it, just run it!
print("Bootstrapped 95% confidence interval for Thai Cafe's true lead over Bua Thai, Pad Thai, and Thai Spoon combined: [{:f}, {:f}]".format(diff_lower_bound, diff_upper_bound))
```

```python
grader.check("q6")
```

## 2. Interpreting Confidence Intervals


We computed the following 95% confidence interval for the percentage of Thai Cafe voters: 

$$[49.40, 54.47]$$

Your answer may have been a bit different; that doesn't mean it was wrong.

<!-- BEGIN QUESTION -->

**Question 7.** Can we say there is a 95% probability that the interval [49.40, 54.47] contains the true percentage of the population that votes for Thai Cafe as the best Durham Thai restaurant? Answer "yes" or "no" and explain your reasoning. 

**Note:** Ambiguous answers using language like "sometimes" or "maybe" will not receive credit.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8.** Suppose your teacher also created 70%, 90%, and 99% confidence intervals from the same sample, but we forgot to label which confidence interval represented which percentages! Match each confidence level (70%, 90%, 99%) with its corresponding interval in the cell below (e.g. % CI: [49.87, 54.0] $\rightarrow$ replace the blank with one of the three confidence levels). Please put them in order from 70% CI, 90% CI, then 99% CI, your teacher thanks you for making this easier to grade! 

**Then**, explain your thought process and how you came up with your answers.

The intervals are below:

* [49.87, 54.00]

* [50.67, 53.27]

* [48.80, 55.40]


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 9.** Suppose you received a grant that allowed you to pay to create 5,000 new samples from the population, each one a uniform random sample of 1,500 Durham residents. If you created a 95% confidence interval from each sample, roughly how many of those 5,000 intervals would you expect to actually contain the true percentage of the population that would vote for Thai Cafe?

Assign your answer to `true_percentage_intervals`.


```python
true_percentage_intervals = ...
```

Recall the second bootstrap confidence interval you created in **Question 6**, which estimated Thai Cafe's lead over Bua Thai, Pad Thai, and Thai Spoon combined. Among voters in the sample, Thai Cafe's lead was 4%. Suppose your teacher's 95% confidence interval for the true lead (in the population of all voters) was

$$[-0.80, 8.80]$$

Suppose we are interested in testing a simple yes-or-no question:

> "Is the percentage of votes for Thai Cafe tied with the percentage of votes for Bua Thai, Pad Thai, and Thai Spoon combined?"

Our null hypothesis is that the percentages are equal, or equivalently, that Thai Cafe's lead is exactly 0. Our alternative hypothesis is that Thai Cafe's lead is not equal to 0.  In the questions below, don't compute any confidence intervals yourself - use only the provided 95% confidence interval.

**Question 10.** Say we use a 5% P-value cutoff.  Do we reject the null, fail to reject the null, or are we unable to tell using the provided confidence interval?

Assign `restaurants_tied` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis

2. Fail to reject the null / Data is consistent with the null hypothesis

3. Unable to tell using our staff confidence interval

**Hint:** If you're confused, take a look at [this chapter](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) of the textbook.


```python
restaurants_tied = ...
```

```python
grader.check("q10")
```

**Question 11.** What if, instead, we use a P-value cutoff of 1%? Do we reject the null, fail to reject the null, or are we unable to tell using the provided confidence interval?

Assign `cutoff_one_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis

2. Fail to reject the null / Data is consistent with the null hypothesis

3. Unable to tell using our staff confidence interval


```python
cutoff_one_percent = ...
```

```python
grader.check("q11")
```

**Question 12.** What if we use a P-value cutoff of 10%? Do we reject, fail to reject, or are we unable to tell using the provided confidence interval?

Assign `cutoff_ten_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis

2. Fail to reject the null / Data is consistent with the null hypothesis

3. Unable to tell using our staff confidence interval


```python
cutoff_ten_percent = ...
```

```python
grader.check("q12")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> type(one_resampled_percentage(votes)) in set([float, np.float64])
True
```
```
>>> 35 <= one_resampled_percentage(votes) <= 65
True
```


Tests for q10 (points: None)
```
>>> restaurants_tied in [1, 2, 3]
True
```


Tests for q11 (points: None)
```
>>> cutoff_one_percent in [1, 2, 3]
True
```


Tests for q12 (points: None)
```
>>> cutoff_ten_percent in [1, 2, 3]
True
```


Tests for q2 (points: None)
```
>>> len(percentages_in_resamples())
2500
```


Tests for q3 (points: None)
```
>>> 40 <= thai_cafe_lower_bound <= thai_cafe_upper_bound <= 60
True
```


Tests for q4 (points: None)
```
>>> type(one_resampled_difference(votes)) in set([float, np.float64])
True
```


Tests for q6 (points: None)
```
>>> -5 <= diff_lower_bound <= diff_upper_bound <= 12
True
```


Tests for q9 (points: None)
