```python
# Initialize Otter
import otter
grader = otter.Notebook("hw06.ipynb")
```

# Run the cell below

To run a code cell (i.e.; execute the python code inside a Jupyter notebook) you can click the play button on the ribbon underneath the name of the notebook. Before you begin click the "Run cell" button at the top that looks like ▶| or hold down `Shift` + `Return`.

# Homework 06: Testing Hypotheses

**Recommended Reading**: 

- [Testing Hypotheses](https://inferentialthinking.com/chapters/11/Testing_Hypotheses.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Notes:** 

- Start early so that you can come to office hours if you're stuck. Check the course website for the office hours schedule. Late work will not be accepted as per the course expectations. Late work will not be accepted without advanced permission as per the course expectations document, so plan accordingly. You should start early so that you have time to get help if you're stuck. Office hours are held throughout the week, so be sure to use them if needed.

- This homework has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. Additional tests will be run once your homework is submitted to Gradescope**. Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and not super far off the expected answer. The tests that are run after submission will evaluate accuracy more carefully. Do not assume that just because all your tests pass means that your answers are correct!

- Directly sharing answers is **not** okay, but discussing problems with the course staff (teacher and TAs) or with other students is encouraged. Refer to the course expectations document to learn more about how to learn and work cooperatively.

```python
# Don't change this cell; just run it. 
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 1. Spam Calls

### 781 Fun

Yanay gets a lot of spam calls. An area code is defined to be a three digit number from 200-999 inclusive. In reality, many of these area codes are not in use, but for this question we'll simplify things and assume they all are. 

**Note:** Throughout these questions, you should assume that Yanay's area code is 781.

**Question 1.** Assuming each area code is just as likely as any other, what's the probability that the area code of two back to back spam calls are 781?


```python
prob_781 = ...
prob_781
```

```python
grader.check("q1")
```

**Question 2.** Rohan already knows that Yanay's area code is 781. Rohan randomly guesses the last 7 digits (0-9 inclusive) of his phone number. What's the probability that Rohan correctly guesses Yanay's number, assuming he’s equally likely to choose any digit?

**Note:** A phone number contains an area code and 7 additional digits, i.e. xxx-xxx-xxxx


```python
prob_yanay_num = ...
prob_yanay_num
```

```python
grader.check("q2")
```

Yanay suspects that there's a higher chance that the spammers are using his area code (781) to trick him into thinking it's someone from his area calling him. Ashley thinks that this is not the case, and that spammers are just choosing area codes of the spam calls at random from all possible area codes (**Remember**, for this question we’re assuming the possible area codes are 200-999, inclusive). Yanay wants to test his claim using the 50 spam calls he received in the past month.

Here's a dataset of the area codes of the 50 spam calls he received in the past month.

```python
# Just run this cell
spam = Table().read_table('data/spam.csv')
spam
```

```python
spam.group('Area Code')
```

<!-- BEGIN QUESTION -->

**Question 3.** Define the null hypothesis and alternative hypothesis for this investigation. 

**Hint:** Don’t forget that your null hypothesis should fully describe a probability model that we can use for simulation later.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 4.** Which of the following test statistics would be a reasonable choice to help differentiate between the two hypotheses?

**Hint:** For a refresher on choosing test statistics, check out the textbook section on [Test Statistics](https://inferentialthinking.com/chapters/11/3/Decisions_and_Uncertainty.html#step-2-the-test-statistic).

1. The proportion of area codes that are 781 in 50 random spam calls

2. The total variation distance (TVD) between probability distribution of randomly chosen area codes, and the observed distribution of area codes.
   
   **Note:** Remember the possible area codes are 200-999 inclusive.

3. The probability of getting an area code of 781 out of all the possible area codes.

4. The proportion of area codes that are 781 in 50 random spam calls divided by 2

5. The number of times you see the area code 781 in 50 random spam calls

Assign `reasonable_test_statistics` to an array of numbers corresponding to these test statistics.


```python
reasonable_test_statistics = ...
```

```python
grader.check("q4")
```

### Please Note:

For the rest of this question, suppose you decide to use **the number of times you see the area code 781 in 50 spam calls** as your test statistic.

**Question 5.** Write a function called `simulate` that generates exactly one simulated value of your test statistic under the null hypothesis.  It should take no arguments and simulate 50 area codes under the assumption that the result of each area is sampled from the range 200-999 inclusive with equal probability. Your function should return the number of times you saw the 781 area code in those 50 random spam calls.

```python
possible_area_codes = ...
def simulate():
    ...
    
# Call your function to make sure it works
simulate()
```

```python
grader.check("q5")
```

**Question 6.** Generate 20,000 simulated values of the number of times you see the area code 781 in 50 random spam calls. Assign `test_statistics_under_null` to an array that stores the result of each of these trials. 

**Hint:** Use the function you defined in **Question 5**.


```python
test_statistics_under_null = ...
repetitions = ...

...
    
test_statistics_under_null
```

```python
grader.check("q6")
```

<!-- BEGIN QUESTION -->

**Question 7.** Using the results from **Question 6**, generate a histogram of the empirical distribution of the number of times you saw the area code 781 in your simulation. 

**NOTE:** Use the provided bins when making the histogram


```python
Table().with_column("Number of 781 area codes (simulated under null)", test_statistics_under_null).group(0).column(1).item(2)/np.sum(test_statistics_under_null)
```

```python
bins = np.arange(0,5,1) # Use these provided bins
...
```

<!-- END QUESTION -->

**Question 8.** Compute an empirical P-value for this test.


```python
# First calculate the observed value of the test statistic from the `spam` table.
observed_val = ...
p_value = ...
p_value
```

```python
grader.check("q8")
```

<!-- BEGIN QUESTION -->

**Question 9.** Suppose you use a P-value cutoff of 1%. What do you conclude from the hypothesis test? Why?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Multiple Spammers

Instead of checking if the area code is equal to his own, Yanay decides to check if the area code matches the area code of one of the 8 places he's been to recently, and wants to test if it's more likely to receive a spam call with an area code from any of those 8 places. These are the area codes of the places he's been to recently: 781, 617, 509, 510, 212, 858, 339, 626.

<!-- BEGIN QUESTION -->

**Question 10.** Define the null hypothesis and alternative hypothesis for this investigation.

**Reminder:** Don’t forget that your null hypothesis should fully describe a probability model that we can use for simulation later.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Please Note:

For the remainder of this part use **the number of times you see any of the area codes of the places Yanay has been to in 50 spam calls** as your test statistic.

**Question 11.** Write a function called `simulate_visited_area_codes` that generates exactly one simulated value of your test statistic under the null hypothesis.  It should take no arguments and simulate 50 area codes under the assumption that the result of each area is sampled from the range 200-999 inclusive with equal probability. Your function should return the number of times you saw any of the area codes of the places Yanay has been to in those 50 spam calls.

**Hint:** You may find the textbook [section](https://inferentialthinking.com/chapters/10/4/Random_Sampling_in_Python.html?highlight=sample_proportion#) on the `sample_proportions` function to be useful.


```python
model_proportions = ...
def simulate_visited_area_codes():
    ...
    
# Call your function to make sure it works
simulate_visited_area_codes()
```

```python
grader.check("q11")
```

**Question 12.** Generate 20,000 simulated values of the number of times you see any of the area codes of the places Yanay has been to in 50 random spam calls. Assign `test_statistics_under_null` to an array that stores the result of each of these trials. 

**Hint:** Use the function you defined in **Question 11**.


```python
visited_test_statistics_under_null = ...

repetitions = ...
...
visited_test_statistics_under_null
```

```python
grader.check("q12")
```

<!-- BEGIN QUESTION -->

**Question 13.** Using the results from **Question 12** generate a histogram of the empirical distribution of the number of times you saw any of the area codes of the places Yanay has been to in your simulation. 

**Note:** Use the provided bins when making the histogram.


```python
bins_visited = np.arange(0,6,1) # Use these provided bins
...
```

<!-- END QUESTION -->

**Question 14.** Compute an empirical P-value for this test.


```python
visited_area_codes = make_array(781, 617, 509, 510, 212, 858, 339, 626)
# First calculate the observed value of the test statistic from the `spam` table.
visited_observed_value = ...
p_value = ...
p_value
```

```python
grader.check("q14")
```

<!-- BEGIN QUESTION -->

**Question 15.** Suppose you use a P-value cutoff of 0.05% 

**Note:** That's 0.05%, not our usual cutoff of 5%. What do you conclude from the hypothesis test? Why?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 16.**

Is `p_value`:

* ( a ) the probability that the spam calls favored the visited area codes,

* ( b ) the probability that they didn't favor, or

* ( c ) neither

Provide a brief justification of your answer. If you chose ( c ), explain what it is instead.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 17.**

Is 0.05% (the P-value cutoff):

* ( a ) the probability that the spam calls favored the visited area codes,

* ( b ) the probability that they didn't favor, or

* ( c ) neither

Provide a brief justification of your answer. If you chose ( c ), explain what it is instead.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 18.** Suppose you run this test for 4000 different people after observing each person's last 50 spam calls. When you reject the null hypothesis for a person, you accuse the spam callers of favoring the area codes that person has visited. If the spam callers were not actually favoring area codes that people have visited, can we compute how many times we will incorrectly accuse the spam callers of favoring area codes that people have visited? If so, what is the number? Explain your answer. Assume a 0.05% P-value cutoff.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. Practice with A/B Tests

Yanay collects information about this month's spam calls. The table `with_labels` is a sampled table, where the `Area Code Visited` column contains either `"Yes"` or `"No"` which represents whether or not Yanay has visited the location of the area code. The `Picked Up` column is `1` if Yanay picked up and `0` if he did not pick up.

```python
# Just run this cell
with_labels = Table().read_table("data/spam_picked_up.csv")
with_labels
```

Yanay is going to perform an A/B Test to see whether or not he is more likely to pick up a call from an area code he has visited. Specifically, his null hypothesis is that there is no difference in the distribution of calls he picked up between visited and not visited area codes, with any difference due to chance. His alternative hypothesis is that there is a difference between the two categories, specifically that he thinks that he is more likely to pick up if he has visited the area code. We are going to perform a [permutation test](https://www.inferentialthinking.com/chapters/12/1/AB_Testing.html#Permutation-Test) to test this. Our test statistic will be the difference in proportion of calls picked up between the area codes Yanay visited and the area codes he did not visit.

**Question 19.** Complete the `difference_in_proportion` function to have it calculate this test statistic, and use it to find the observed value. The function takes in a sampled table which can be any table that has the same columns as `with_labels`. We'll call `difference_in_proportion` with the sampled table `with_labels` in order to find the observed difference in proportion.


```python
def difference_in_proportion(sample):
    # Take a look at the code for `proportion_visited` and use that as a 
    # hint of what `proportions` should be assigned to
    proportions = ...
    proportion_visited = proportions.where("Area Code Visited", "Yes").column("Picked Up mean").item(0)
    proportion_not_visited = proportions.where("Area Code Visited", "No").column("Picked Up mean").item(0)
    ...

observed_diff_proportion = difference_in_proportion(with_labels)
observed_diff_proportion
```

```python
grader.check("q19")
```

**Question 20.** To perform a permutation test we shuffle the labels, because our null hypothesis is that the labels don't matter because the distribution of calls he picked up between visited and not visited area codes come from same underlying distribution. The labels in this case is the `"Area Code Visited"` column containing `"Yes"` and `"No"`.

Write a function to shuffle the table and return a test statistic using the function you defined in **Question 19**.

**Hint:** To shuffle labels, you can sample without replacement and then replace the appropriate column with the new shuffled column.


```python
def simulate_one_stat():
    shuffled = ...
    original_with_shuffled_labels = ...
    return difference_in_proportion(original_with_shuffled_labels)

one_simulated_test_stat = simulate_one_stat() 
one_simulated_test_stat
```

```python
grader.check("q20")
```

<!-- BEGIN QUESTION -->

**Question 21.** Generate 1,000 simulated test statistic values. Assign `test_stats` to an array that stores the result of each of these trials. 

**Hint**: Use the function you defined in **Question 20**.

We also provided code that'll generate a histogram for you after generating a 1000 simulated test statistic values.


```python
trials = ...
test_stats = ...

...

# here's code to generate a histogram of values and the red dot is the observed value
Table().with_column("Simulated Proportion Difference", test_stats).hist("Simulated Proportion Difference");
plt.plot(observed_diff_proportion, 0, 'ro', markersize=15);
```

<!-- END QUESTION -->

**Question 22.** Compute the empirical p-value for this test, and assign it to `p_value_ab`.


```python
p_value_ab = ...
p_value_ab
```

```python
grader.check("q22")
```

For `p_value_ab`, you should be getting a value around 10-15%. If our p-value cutoff is 5%, the data is more consistent with the null hypothesis - that there is no difference in the distribution of calls Yanay picked up between visited and not visited area codes.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> 0 <= prob_781 <= 1
True
```


Tests for q11 (points: None)
```
>>> np.std([simulate_visited_area_codes() for _ in range(1000)]) > 0
True
```
```
>>> sum(model_proportions) == 1
True
```


Tests for q12 (points: None)
```
>>> len(visited_test_statistics_under_null) == 20000
True
```


Tests for q14 (points: None)
```
>>> 0 <= p_value < 1
True
```


Tests for q19 (points: None)
```
>>> -1 <= observed_diff_proportion <= 1
True
```
```
>>> np.round(observed_diff_proportion, 3) == 0.219
True
```


Tests for q2 (points: None)
```
>>> 0 <= prob_yanay_num <= 1
True
```


Tests for q20 (points: None)
```
>>> -0.75 <= one_simulated_test_stat <= 0.75
True
```


Tests for q22 (points: None)
```
>>> p_value_ab > 0.05
True
```


Tests for q4 (points: None)
```
>>> isinstance(reasonable_test_statistics, np.ndarray)
True
```
```
>>> len(reasonable_test_statistics) > 0
True
```


Tests for q5 (points: None)
```
>>> np.std([simulate() for _ in range(1000)]) > 0
True
```


Tests for q6 (points: None)
```
>>> len(test_statistics_under_null) == 20000
True
```


Tests for q8 (points: None)
```
>>> 0 <= p_value < 1
True
```
