```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# Lab 02: Table Operations

Welcome to Lab 02! 

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Collaboration Policy:**

Collaborating on labs is more than okay -- it's encouraged! You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** _just_ copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

To receive credit for this assignment, answer all questions correctly and submit before the deadline.


# Today's Lab 

This week, we'll learn how to import a module and practice table operations.

Recommended reading:

 * [Names](https://inferentialthinking.com/chapters/03/2/Names.html)
 
 * [Call Expressions](https://inferentialthinking.com/chapters/03/3/Calls.html)
 
 * [Introduction to Tables](https://www.inferentialthinking.com/chapters/03/4/Introduction_to_Tables)

First, set up the imports.

```python
from datascience import *
```

# 1. Review: Introduction to Python

The two building blocks of Python code are *expressions* and *statements*.  An **expression** is a piece of code that

* is self-contained, meaning it would make sense to write it on a line by itself, and
* usually evaluates to a value.


Here are two expressions that both evaluate to 3

* 3
* 5 - 2
    
One important type of expression is the **call expression**. A call expression begins with the name of a function and is followed by the argument(s) of that function in parentheses. The function returns some value, based on its arguments. Some important mathematical functions are listed below.

| Function | Description                                                   |
|----------|---------------------------------------------------------------|
| `abs`      | Returns the absolute value of its argument                    |
| `max`      | Returns the maximum of all its arguments                      |
| `min`      | Returns the minimum of all its arguments                      |
| `pow`      | Raises its first argument to the power of its second argument |
| `round`    | Rounds its argument to the nearest integer                     |

Here are two call expressions that both evaluate to 3

* `abs(2 - 5)`

* `max(round(2.8), min(pow(2, 10), -1 * pow(2, 10)))`

The expression `5 - 2` and the two call expressions given above are examples of **compound expressions**, meaning that they are actually combinations of several smaller expressions.  `5 - 2` combines the expressions `5` and `2` by subtraction.  In this case, `5` and `2` are called **subexpressions** because they're expressions that are part of a larger expression.

A **statement** is a whole line of code.  Some statements are just expressions.  The expressions listed above are examples.

Other statements *make something happen* rather than *having a value*. For example, an **assignment statement** assigns a value to a name. 

A good way to think about this is that we're **evaluating the right-hand side** of the equals sign and **assigning it to the left-hand side**. Here are some assignment statements:
    
* `height = 1.3`

* `the_number_five = abs(-5)`

* `absolute_height_difference = abs(height - 1.688)`

An important idea in programming is that large, interesting things can be built by combining many simple, uninteresting things.  The key to understanding a complicated piece of code is breaking it down into its simple components.

For example, a lot is going on in the last statement above, but it's really just a combination of a few things.  This picture describes what's going on.

<img src="images/statement.png"/>

**Question 1.** In the next cell, assign the name `new_year` to the larger number among the following two numbers:

1. the **absolute value** of $2^{5}-2^{11}-2^2 -4$ , and 

1. $5 \times 13 \times 31+2$.

Try to use just one statement (one line of code). Be sure to check your work by executing the test cell afterward.


```python
new_year = ...
new_year
```

```python
grader.check("q1")
```

We've asked you to use one line of code in the question above because it only involves mathematical operations. However, more complicated programming questions will more require more steps. It isn’t always a good idea to jam these steps into a single line because it can make the code harder to read and harder to debug.

Good programming practice involves splitting up your code into smaller steps and using appropriate names. You'll have plenty of opportunities to practice this throughout this course.

# 2. Importing code

> What has been will be again,  
> what has been done will be done again;  
> there is nothing new under the sun.

Most programming involves work that is very similar to work that has been done before.  Since writing code is time-consuming, it's good to rely on others' published code when you can.  Rather than copy-pasting, Python allows us to **import modules**. A module is a file with Python code that has defined variables and functions. By importing a module, we are able to use its code in our own notebook.

Python includes many useful modules that are just an `import` away.  We'll look at the `math` module as a first example. The `math` module is extremely useful in computing mathematical expressions in Python. 

Suppose we want to very accurately compute the area of a circle with a radius of 5 meters.  For that, we need the constant $\pi$, which is roughly 3.14.  Conveniently, the `math` module has `pi` defined for us.

Run the cell below.

```python
import math

radius = 5
area_of_circle = radius**2 * math.pi
area_of_circle
```

In the code above, the line `import math` imports the math module. This statement creates a module and then assigns the name `math` to that module. We are now able to access any variables or functions defined within `math` by typing the name of the module followed by a dot, then followed by the name of the variable or function we want. 

For example,

`<module name>.<name>`

**Question 2.** The module `math` also provides the name `e` for the base of the natural logarithm, which is roughly 2.71.  Compute $e^{\pi}-\pi$, giving it the name `near_twenty`.


```python
near_twenty = ...
near_twenty
```

```python
grader.check("q2")
```

## 2.1. Accessing Functions

In the question above, you accessed variables within the `math` module. 

**Modules** also define **functions**.  For example, `math` provides the name `sin` for the sine function.  Having imported `math` already, we can write `math.sin(3)` to compute the sine of 3.  (Note that this sine function considers its argument to be in [radians](https://en.wikipedia.org/wiki/Radian), not degrees.  180 degrees are equivalent to $\pi$ radians.)

**Question 3.** A $\frac{\pi}{4}$ radian (45-degree) angle forms a right triangle with equal base and height, pictured below.  If the hypotenuse (the radius of the circle in the picture) is 1, then the height is $\sin\left(\frac{\pi}{4}\right)$.  Compute that value using `sin` and `pi` from the `math` module.  Give the result the name `sine_of_pi_over_four`.

<img src="http://mathworld.wolfram.com/images/eps-gif/TrigonometryAnglesPi4_1000.gif">

**Source:** [Wolfram MathWorld](http://mathworld.wolfram.com/images/eps-gif/TrigonometryAnglesPi4_1000.gif)



```python
sine_of_pi_over_four = ...
sine_of_pi_over_four
```

```python
grader.check("q3")
```

For your reference, below are some more examples of functions from the `math` module.

Notice how different functions take in different numbers of arguments. Often, the [documentation](https://docs.python.org/3/library/math.html) of the module will provide information on how many arguments are required for each function.

Run the cell below.

```python
# Calculating logarithms (the logarithm of 8 in base 2)
# The result is 3 because 2 to the power of 3 is 8
math.log(8, 2)
```

Run the cell below.

```python
# Calculating square roots
math.sqrt(5)
```

There are various ways to import and access code from outside sources. The method we used above — `import <module_name>` — imports the entire module and requires that we use `<module_name>.<name>` to access its code. 

We can also import a specific constant or function instead of the entire module. Notice that you don't have to use the module name beforehand to reference that particular value. However, you do have to be careful about reassigning the names of the constants or functions to other values.

Run the cell below.

```python
# Importing just cos and pi from math
# We don't have to use `math.` in front of cos or pi

from math import cos, pi
print(cos(pi))

# We do have to use it in front of other functions from math, though
math.log(pi)
```

Or we can import every function and value from the entire module.

Run the cell below.

```python
# Lastly, we can import everything from math using the *
# Once again, we don't have to use 'math.' beforehand 

from math import *
log(pi)
```

Don't worry too much about which type of import to use. It's often a coding style choice left up to each programmer. In this course, you'll always import the necessary modules when you run the setup cell (like the first code cell in this lab).

# 3. Table operations

The table `farmers_markets.csv` contains data on farmers' markets in the United States  (data collected [by the USDA](https://apps.ams.usda.gov/FarmersMarketsExport/ExcelExport.aspx)).  Each row represents one such market.

Run the cell below.

```python
farmers_markets = Table.read_table('data/farmers_markets.csv')
```

<!-- BEGIN QUESTION -->

Let's examine our table to see what data it contains.

**Question 4.** Use the method `show` to display the first 5 rows of `farmers_markets`.

**Note:** The terms "method" and "function" are technically not the same thing, but for the purposes of this course, we will use them interchangeably.

**Warning:** Make sure not to call `.show()` without an argument, as this will **crash your kernel**.


```python
...
```

<!-- END QUESTION -->

Notice that some of the values in this table are missing, as denoted by "nan." This means either that the value is not available (e.g. if we don’t know the market’s street address) or not applicable (e.g. if the market doesn’t have a street address). You'll also notice that the table has a large number of columns in it.

The table property `num_columns` returns the number of columns in a table. A "property" is just a method that doesn't need to be called by adding parentheses.

Example call: `<tbl>.num_columns`

**Question 5.** Use `num_columns` to find the number of columns in our farmers' markets dataset.

Assign the number of columns to `num_farmers_markets_columns`.


```python
num_farmers_markets_columns = ...
print('The table has', num_farmers_markets_columns, 'columns in it.')
```

```python
grader.check("q5")
```

Similarly, the property `num_rows` tells you how many rows are in a table.

**Question 6.** Assign the number of rows to `num_farmers_markets_rows`.


```python
num_farmers_markets_rows = ...
print("The table has", num_farmers_markets_rows, "rows in it.")
```

```python
grader.check("q6")
```

Most of the columns are about particular products -- whether the market sells tofu, pet food, etc.  If we're not interested in that information, it just makes the table difficult to read.  This comes up more than you might think, because people who collect and publish data may not know ahead of time what people will want to do with it.

In such situations, we can use the table method `select` to choose only the columns that we want in a particular table. It takes any number of arguments. Each should be the name of a column in the table. It returns a new table with only those columns in it. The columns are in the order in which they were listed as arguments.

For example, the value of `farmers_markets.select("MarketName", "State")` is a table with only the name and the state of each farmers' market in `farmers_markets`.



**Question 7.** Use `select` to create a table with only the name, city, state, latitude (`y`), and longitude (`x`) of each market.  Call that new table `farmers_markets_locations`.

**Hint:** Make sure to be exact when using column names with `select`; double-check capitalization.


```python
farmers_markets_locations = ...
farmers_markets_locations
```

```python
grader.check("q7")
```

`drop` serves the same purpose as `select`, but it takes away the columns that you provide rather than the ones that you don't provide. Like `select`, `drop` returns a new table.

**Question 8.** Suppose you just didn't want the `FMID` or `updateTime` columns in `farmers_markets`.  Create a table that's a copy of `farmers_markets` but doesn't include those columns.  Call that table `farmers_markets_without_fmid`.


```python
farmers_markets_without_fmid = ...
```

```python
grader.check("q8")
```

Now, suppose we want to answer some questions about farmers' markets in the US. For example, which market(s) have the largest longitude (given by the `x` column)? 

To answer this, we'll sort `farmers_markets_locations` by longitude.

Run the cell below.

```python
farmers_markets_locations.sort('x')
```

That didn't answer our question because we sorted from smallest to largest longitude. To look at the largest longitudes, we'll have to sort in reverse order.

**Question 9.** Sort the `farmers_markets_locations` table from largest to smallest longitude. Call it `farmers_markets_locations_by_longitude`.


```python
farmers_markets_locations_by_longitude = ...
farmers_markets_locations_by_longitude
```

```python
grader.check("q9")
```

Some details about sort

1. The first argument to `sort` is the name of a column to sort by.

1. If the column has text in it, `sort` will sort alphabetically; if the column has numbers, it will sort numerically.

1. The value of `farmers_markets_locations.sort("x")` is a *copy* of `farmers_markets_locations`; the `farmers_markets_locations` table doesn't get modified. For example, if we called `farmers_markets_locations.sort("x")`, then running `farmers_markets_locations` by itself would still return the unsorted table.

1. Rows always stick together when a table is sorted.  It wouldn't make sense to sort just one column and leave the other columns alone.  For example, in this case, if we sorted just the `x` column, the farmers' markets would all end up with the wrong longitudes.

**Question 10.** Create a version of `farmers_markets_locations` that's sorted by **latitude (`y`)**, with the largest latitudes first.  Call it `farmers_markets_locations_by_latitude`.


```python
farmers_markets_locations_by_latitude = ...
farmers_markets_locations_by_latitude
```

```python
grader.check("q10")
```

Now let's say we want a table of all farmers' markets in North Carolina. Sorting won't help us much here because North Carolina is closer to the middle of the dataset.

Instead, we use the table method `where`. Look at the [documentation](http://data8.org/datascience/) to see how to use `where`.

Run the following cell:

```python
nc_farmers_markets = farmers_markets_locations.where('State', are.equal_to('North Carolina'))
nc_farmers_markets
```

Ignore the syntax for the moment.  Instead, try to read that line like this:

> Assign the name **`nc_farmers_markets`** to a table whose rows are the rows in the **`farmers_markets_locations`** table **`where`** the `State`'s **`are` `equal` `to` `North Carolina`**.

Now let's dive into the details a bit more.

`where` takes 2 arguments

1. The name of a column.  `where` finds rows where that column's values meet some criterion.

1. A predicate that describes the criterion that the column needs to meet.

The predicate in the example above called the function `are.equal_to` with the value we wanted, 'North Carolina'.  We'll see other predicates soon.

`where` returns a table that's a copy of the original table, but **with only the rows that meet the given predicate**.

**Question 11.** Use `nc_farmers_markets` to create a table called `durham_farmers_markets` containing farmers' markets in Durham, North Carolina.


```python
durham_farmers_markets = ...
durham_farmers_markets
```

```python
grader.check("q11")
```

So far we've only been using `where` with the predicate that requires finding the values in a column to be *exactly* equal to a certain value. However, there are many other predicates. Here are a few:

|Predicate|Example|Result|
|-|-|-|
|`are.equal_to`|`are.equal_to(50)`|Find rows with values equal to 50|
|`are.not_equal_to`|`are.not_equal_to(50)`|Find rows with values not equal to 50|
|`are.above`|`are.above(50)`|Find rows with values above (and not equal to) 50|
|`are.above_or_equal_to`|`are.above_or_equal_to(50)`|Find rows with values above 50 or equal to 50|
|`are.below`|`are.below(50)`|Find rows with values below 50|
|`are.between`|`are.between(2, 10)`|Find rows with values above or equal to 2 and below 10|

## 4. Analyzing a dataset
Now that you're familiar with table operations, let’s answer some interesting questions about a dataset. Run the cell below to load the `imdb` table. It contains information about the 250 highest-rated movies on IMDb.

```python
imdb = Table.read_table('data/imdb.csv')
imdb
```

Often, we want to perform multiple operations - sorting, filtering, or others - in order to turn a table we have into something more useful. You can do these operations one by one, e.g.

```
first_step = original_tbl.where(“col1”, are.equal_to(12))
second_step = first_step.sort(‘col2’, descending=True)
```

However, since the value of the expression `original_tbl.where(“col1”, are.equal_to(12))` is itself a table, you can just call a table method on it:

```
original_tbl.where(“col1”, are.equal_to(12)).sort(‘col2’, descending=True)
```
You should organize your work in the way that makes the most sense to you, using informative names for any intermediate tables you create.

**Question 12.** Create a table of movies released between 2010 and 2016 (inclusive) with ratings above 8. The table should only contain the columns `Title` and `Rating`, **in that order**.

Assign the table to the name `above_eight`.

**Hint:** Think about the steps you need to take, and try to put them in an order that make sense. Feel free to create intermediate tables for each step, but please make sure you assign your final table the name `above_eight`!

```python
above_eight = ...
above_eight
```

```python
grader.check("q12")
```

**Question 13.** Use `num_rows` (and arithmetic) to find the *proportion* of movies in the dataset that were released 1900-1999, and the *proportion* of movies in the dataset that were released in the year 2000 or later.

Assign `proportion_in_20th_century` to the proportion of movies in the dataset that were released 1900-1999, and `proportion_in_21st_century` to the proportion of movies in the dataset that were released in the year 2000 or later.

**Hint:** The *proportion* of movies released in the 1900's is the *number* of movies released in the 1900's, divided by the *total number* of movies.


```python
proportion_in_20th_century = ...
proportion_in_21st_century = ...
print('20th Century: ', proportion_in_20th_century)
print('21st Century: ', proportion_in_21st_century)
```

```python
grader.check("q13")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by `SHIFT`-clicking on the file name and selecting **Save Link As**. Or, find the .zip file in the left side of the screen and right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> new_year != ...
True
```
```
>>> new_year != 2016
True
```
```
>>> new_year == 2024
True
```


Tests for q10 (points: None)
```
>>> farmers_markets_locations_by_latitude.first('y') == 64.86275
True
```


Tests for q11 (points: None)
```
>>> durham_farmers_markets.num_rows == 5
True
```
```
>>> 'Durham' in durham_farmers_markets['city']
True
```


Tests for q12 (points: None)
```
>>> above_eight.num_rows == 28
True
```
```
>>> above_eight.sort(0).take([17])[0][0] == 'Shutter Island'
True
```


Tests for q13 (points: None)
```
>>> proportion_in_20th_century == 0.64
True
```
```
>>> proportion_in_21st_century == 0.36
True
```


Tests for q2 (points: None)
```
>>> near_twenty != math.e ** math.pi
True
```
```
>>> near_twenty == math.e ** math.pi - math.pi
True
```


Tests for q3 (points: None)
```
>>> sine_of_pi_over_four == math.sin(math.pi / 4)
True
```


Tests for q5 (points: None)
```
>>> num_farmers_markets_columns == farmers_markets.num_columns
True
```


Tests for q6 (points: None)
```
>>> num_farmers_markets_rows == farmers_markets.num_rows
True
```


Tests for q7 (points: None)
```
>>> farmers_markets_locations.num_columns == 5
True
```
```
>>> sorted(farmers_markets_locations.labels) == ['MarketName', 'State', 'city', 'x', 'y']
True
```
```
>>> farmers_markets_locations.num_rows == 8546
True
```


Tests for q8 (points: None)
```
>>> farmers_markets_without_fmid.num_columns == 57
True
```
```
>>> sorted(farmers_markets_without_fmid.labels) == ['Bakedgoods', 'Beans', 'Cheese', 'Coffee', 'County', 'Crafts', 'Credit', 'Eggs', 'Facebook', 'Flowers', 'Fruits', 'Grains', 'Herbs', 'Honey', 'Jams', 'Juices', 'Location', 'Maple', 'MarketName', 'Meat', 'Mushrooms', 'Nursery', 'Nuts', 'Organic', 'OtherMedia', 'PetFood', 'Plants', 'Poultry', 'Prepared', 'SFMNP', 'SNAP', 'Seafood', 'Season1Date', 'Season1Time', 'Season2Date', 'Season2Time', 'Season3Date', 'Season3Time', 'Season4Date', 'Season4Time', 'Soap', 'State', 'Tofu', 'Trees', 'Twitter', 'Vegetables', 'WIC', 'WICcash', 'Website', 'WildHarvested', 'Wine', 'Youtube', 'city', 'street', 'x', 'y', 'zip']
True
```


Tests for q9 (points: None)
```
>>> type(farmers_markets_locations_by_longitude) == tables.Table
True
```
```
>>> list(farmers_markets_locations_by_longitude.column('x').take(range(3))) == [-64.7043, -64.7789, -64.8799]
True
```
