```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

# Homework 04: Functions, Histograms and Groups

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided modules. Each time you start your server, you will need to execute this cell again to load the modules.

```python
# Don't change this cell; just run it. 
import numpy as np
from datascience import *

import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

**Recommended Reading:** 

* [Visualizing Numerical Distributions](https://inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html) 

* [Functions and Tables](https://inferentialthinking.com/chapters/08/Functions_and_Tables.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Notes:** 

- Start early so that you can come to office hours if you're stuck. Check the course website for the office hours schedule. Late work will not be accepted as per the course expectations. Late work will not be accepted without advanced permission as per the course expectations document, so plan accordingly. You should start early so that you have time to get help if you're stuck. Office hours are held throughout the week, so be sure to use them if needed.

- This homework has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. Additional tests will be run once your homework is submitted to Gradescope**. Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and not super far off the expected answer. The tests that are run after submission will evaluate accuracy more carefully. Do not assume that just because all your tests pass means that your answers are correct!

- Directly sharing answers is **not** okay, but discussing problems with the course staff (teacher and TAs) or with other students is encouraged. Refer to the course expectations document to learn more about how to learn and work cooperatively.

## 1. Burrito-ful San Diego

Tam, Margaret and Winifred are trying to use Data Science to find the best burritos in San Diego! Their friends Irene and Maya provided them with two comprehensive datasets on many burrito establishments in the San Diego area taken from (and cleaned from): https://www.kaggle.com/srcole/burritos-in-san-diego/data.

The following cell reads in a table called `ratings` which contains names of burrito restaurants, their Yelp rating, Google rating, as well as their Overall rating. It also reads in a table called `burritos_types` which contains names of burrito restaurants, their menu items, and the cost of the respective menu item at the restaurant.

```python
# Just run this cell
ratings = Table.read_table("data/ratings.csv")
ratings.show(5)
burritos_types = Table.read_table("data/burritos_types.csv")
burritos_types.show(5)
```

**Question 1.** It would be easier if we could combine the information in both tables. Assign `burritos` to the result of joining the two tables together.

**Note:** it doesn't matter which table you put in as the argument to the table method, either order will work for the autograder tests.

**Hint:** If you need refreshers on table methods, look at the [python reference](http://data8.org/sp20/python-reference.html).

```python
burritos = ...
burritos.show(5)
```

```python
grader.check("q1")
```

<!-- BEGIN QUESTION -->

**Question 2.** Let's look at how the Yelp scores compare to the Google scores in the `burritos` table. First, assign `yelp_and_google` to a table only containing the columns `Yelp` and `Google`. Then, make a scatter plot with Yelp scores on the x-axis and the Google scores on the $y-$axis. 


```python
yelp_and_google = ...
...
# Don't change/edit/remove the following line.
# To help you make conclusions, we have plotted a straight line on the graph (y=x)
plt.plot(np.arange(2.5,5,.5), np.arange(2.5,5,.5));
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.** Looking at the scatter plot you just made in **Question 2.**, do you notice any pattern(s) (i.e. is one of the two types of scores consistently higher than the other one, is there any association between the two)? If so, describe them **briefly** in the cell below.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

Here's a refresher on how `.group` works! You can read how `.group` works in the [textbook](https://www.inferentialthinking.com/chapters/08/2/Classifying_by_One_Variable.html), or you can view the video below. 

```python
from IPython.display import YouTubeVideo
YouTubeVideo("HLoYTCUP0fc")
```

**Question 4.** From the `burritos` table, some of the restaurant locations have multiple reviews. Winifred thinks California burritos are the best type of burritos, and wants to see the average overall rating for California burritos at each location. Create a table that has two columns: the name of the restaurant and the average overall rating of California burritos at each location.

**Tip:** Revisit the burritos table to see how California burritos are represented.

**Note:** You can break up the solution into multiple lines, as long as you assign the final output table to `california_burritos`! For reference however, the instructor solution only used one line.


```python
california_burritos = ...
california_burritos
```

```python
grader.check("q4")
```

**Question 5.** Given this new table `california_burritos`, Winifred can figure out the name of the restaurant with the highest overall average rating! Assign `best_restaurant` to a line of code that evaluates to a string that corresponds to the name of the restaurant with the highest overall average rating. 


```python
best_restaurant = ...
best_restaurant
```

```python
grader.check("q5")
```

<!-- BEGIN QUESTION -->

**Question 6.** Tam thinks that burritos in San Diego are cheaper (and taste better) than the burritos in Berkeley. Plot a histogram that visualizes that distribution of the costs of the burritos from San Diego in the `burritos` table. Also use the provided `bins` variable when making your histogram, so that visually the histogram is more informative.


```python
my_bins = np.arange(0, 15, 1)
# Please also use the provided bins
...
```

<!-- END QUESTION -->

## 2. Faculty Salaries

This exercise is designed to give you practice using the Table methods `pivot` and `group`. [Here](http://data8.org/sp21/python-reference.html) is a link to the Python reference page in case you need a quick refresher.

Run the cell below to view a demo on how you can use pivot on a table.

```python
from IPython.display import YouTubeVideo
YouTubeVideo("4WzXo8eKLAg")
```

In the next cell, we load a dataset maintained by the University of North Carolina System Office which contains information on UNC System faculty including their institution, age, job category, and base salary as of December 2020.

```python
# Just run this cell
unc = Table.read_table("data/unc_faculty.csv").drop('INITIAL HIRE DATE', 'EMPLOYEE HOME DEPARTMENT', 'PRIMARY WORKING TITLE').relabeled('INSTITUTION NAME', 'School Name').relabeled('AGE', 'Age').relabeled('JOB CATEGORY', 'Title').relabeled('EMPLOYEE ANNUAL BASE SALARY', 'Salary')
unc
```

**Question 7.** Set `average_unc_stats` to a table with 4 columns that shows the average age and salary for each combination of `School Name` and `Title` found in the `unc` table. For example, you should be able to determine the average age and salary for an Assistant Professor at ASU by looking up the corresponding row in the table you create.

**Hint:** Use the `.group` method to create this table. You may get a warning about creating an array from "ragged nested sequences". It is safe to ignore.


```python
average_unc_stats = ...
average_unc_stats
```

```python
grader.check("q7")
```

**Question 8.** Now create a pivot table assigned to `unc_pivot` that has each `School Name` for the column labels and a row that corresponds to each unique `Title`. The values in each column should be the average age for each `Title` at that school.


```python
unc_pivot = ...
unc_pivot
```

```python
grader.check("q8")
```

**Question 9.** Write your own function named `salary_range` that can take in an array of salaries and returns the range (the smallest salary subtracted from the the largest salary) in that array.


```python
def salary_range(...):
    ...
```

```python
grader.check("q9")
```

**Question 10.** Set `job_ranges` to a table containing Titles as the rows, and the School as the columns. The values in the rows should correspond to the salary range for the job, where the range is calculated using the `salary_range` function you wrote in **Question 9.**


```python
job_ranges = ...
job_ranges
```

```python
grader.check("q10")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> set(['Name', 'Menu_Item', 'Yelp', 'Google', 'Overall', 'Cost']) == set(burritos.labels)
True
```
```
>>> burritos.num_rows
212
```


Tests for q10 (points: None)
```
>>> job_ranges.num_columns
17
```
```
>>> job_ranges.num_rows
6
```


Tests for q4 (points: None)
```
>>> california_burritos.num_columns
2
```
```
>>> california_burritos.num_rows
46
```


Tests for q5 (points: None)
```
>>> isinstance(best_restaurant, str)
True
```


Tests for q7 (points: None)
```
>>> average_unc_stats.num_columns
4
```
```
>>> average_unc_stats.num_rows
74
```
```
>>> set(average_unc_stats.labels) == set(('School Name', 'Title', 'Age average', 'Salary average'))
True
```
```
>>> np.isclose(sum(average_unc_stats.column('Age average')), 3919.262733622144)
True
```
```
>>> np.isclose(sum(average_unc_stats.column('Salary average')), 7382472.347677045)
True
```


Tests for q8 (points: None)
```
>>> unc_pivot.num_columns
17
```
```
>>> unc_pivot.num_rows
6
```
```
>>> round(unc_pivot.column('WSSU').take(0), 2)
47.57
```


Tests for q9 (points: None)
```
>>> salary_range(make_array(5, 1, 20, 1000))
999
```
