```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

# Lab 05: Simulations

Welcome to Lab 05!


We will go over [iteration](https://www.inferentialthinking.com/chapters/09/2/Iteration.html) and [simulations](https://www.inferentialthinking.com/chapters/09/3/Simulation.html), as well as introduce the concept of [randomness](https://www.inferentialthinking.com/chapters/09/Randomness.html).

The data used in this lab will contain salary data and other statistics for basketball players from the 2014-2015 NBA season. This data was collected from the following sports analytic sites: [Basketball Reference](http://www.basketball-reference.com) and [Spotrac](http://www.spotrac.com).


**Collaboration Policy:** Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here. 

First, set up the notebook by running the cell below.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 1. Nachos and Conditionals

In Python, the boolean data type contains only two unique values:  `True` and `False`. Expressions containing comparison operators such as `<` (less than), `>` (greater than), and `==` (equal to) evaluate to Boolean values. A list of common comparison operators can be found below!

<img src="images/comparisons.png">

Run the cell below to see an example of a comparison operator in action.

```python
3 > 1 + 1
```

We can even assign the result of a comparison operation to a variable.

```python
result = 10 / 2 == 5
result
```

Arrays are compatible with comparison operators. The output is an array of boolean values.

```python
make_array(1, 5, 7, 8, 3, -1) > 3
```

One day, when you come home after a long week, you see a hot bowl of nachos waiting on the dining table! Let's say that whenever you take a nacho from the bowl, it will either have only **cheese**, only **salsa**, **both** cheese and salsa, or **neither** cheese nor salsa (a sad tortilla chip indeed). 

Let's try and simulate taking nachos from the bowl at random using the function, `np.random.choice(...)`.

### `np.random.choice`

`np.random.choice` picks one item at random from the given array. It is equally likely to pick any of the items. Run the cell below several times, and observe how the results change.

```python
nachos = make_array('cheese', 'salsa', 'both', 'neither')
np.random.choice(nachos)
```

To repeat this process multiple times, pass in an int `n` as the second argument to return `n` different random choices. By default, `np.random.choice` samples **with replacement** and returns an *array* of items. 

Run the next cell to see an example of sampling with replacement 10 times from the `nachos` array.

```python
np.random.choice(nachos, 10)
```

To count the number of times a certain type of nacho is randomly chosen, we can use `np.count_nonzero`.

### `np.count_nonzero`

`np.count_nonzero` counts the number of non-zero values that appear in an array. When an array of boolean values are passed through the function, it will count the number of `True` values (remember that in Python, `True` is coded as 1 and `False` is coded as 0.)

Run the next cell to see an example that uses `np.count_nonzero`.

```python
np.count_nonzero(make_array(True, False, False, True, True))
```

**Question 1.** Assume we took ten nachos at random, and stored the results in an array called `ten_nachos` as done below. Find the number of nachos with only cheese using code (do not hardcode the answer).  

**Note:** To earn all the points for this question you must use the comparison operator (e.g. `=`, `<`, ...) and the `np.count_nonzero` method.

```python
ten_nachos = make_array('neither', 'cheese', 'both', 'both', 'cheese', 'salsa', 'both', 'neither', 'cheese', 'both')
number_cheese = ...
number_cheese
```

```python
grader.check("q1")
```

### Conditional Statements

A conditional statement is a multi-line statement that allows Python to choose among different alternatives based on the truth value of an expression.

Here is a basic example.

```
def sign(x):
    if x > 0:
        return 'Positive'
    else:
        return 'Negative'
```

If the input `x` is greater than `0`, we return the string `'Positive'`. Otherwise, we return `'Negative'`.

If we want to test multiple conditions at once, we use the following general format.

```
if <if expression>:
    <if body>
elif <elif expression 0>:
    <elif body 0>
elif <elif expression 1>:
    <elif body 1>
...
else:
    <else body>
```

Only the body for the first conditional expression that is true will be evaluated. Each `if` and `elif` expression is evaluated and considered in order, starting at the top. As soon as a true value is found, the corresponding body is executed, and the rest of the conditional statement is skipped. If none of the `if` or `elif` expressions are true, then the `else body` is executed. 

For more examples and explanation, refer to the section on conditional statements [here](https://www.inferentialthinking.com/chapters/09/1/conditional-statements.html).

**Question 2.** Complete the following conditional statement so that the string `'More please'` is assigned to the variable `say_please` if the number of nachos with cheese in `ten_nachos` is less than `5`.

**Hint:** You could use `number_cheese` from **Question 1.**, since it already contains the number of nachos with cheese in the array `ten_nachos`.

```python
say_please = '?'
if ...:
    say_please = 'More please'
say_please
```

```python
grader.check("q2")
```

**Question 3.** Write a function called `nacho_reaction` that can be provided a nacho type (as a string) and returns a reaction (as a string) based on the type of nacho passed in as an argument. Use the table below to match the nacho type to the appropriate reaction.

| Nacho Type | Reaction |
|------------|----------|
| cheese     | Cheesy!  |
| salsa      | Spicy!   |
| both       | Wow!     |
| neither    | Meh.     |

**Hint:** If you're failing the test, double check the spelling of your reactions.

```python
def nacho_reaction(nacho):
    if nacho == "cheese":
        return ...
    ... :
        ...
    ... :
        ...
    ... :
        ...

spicy_nacho = nacho_reaction('salsa')
spicy_nacho
```

```python
grader.check("q3")
```

**Question 4.** Create a Table named `ten_nachos_reactions` that consists of the nachos in `ten_nachos` as well as the reactions for each of those nachos. The columns should be called `Nachos` and `Reactions`. The starter code will create the Table `ten_nachos_tbl` that already has the column `Nachos` provided. 

**Hint:** Use the `apply` method to create the array for second column. Here's a link to the textbook section if you need a refresher click [here](https://www.inferentialthinking.com/chapters/08/1/Applying_a_Function_to_a_Column.html ).

```python
ten_nachos_tbl = Table().with_column('Nachos', ten_nachos)
...
ten_nachos_reactions
```

```python
grader.check("q4")
```

**Question 5.** Using code, find the number of 'Wow!' reactions for the nachos in `ten_nachos_reactions`. Your code should reference the table `ten_nachos_reactions` at least once.

**Hint:** There are a few ways to do this! Try seeing if you can use a conditional statement on an array to help out.

```python
number_wow_reactions = ...
number_wow_reactions
```

```python
grader.check("q5")
```

## 2. Simulations and For Loops
Using a `for` statement, we can perform a task multiple times. This is known as iteration.

One use of iteration is to loop through a set of values. For instance, we can print out all of the colors of the rainbow.

```python
rainbow = make_array("red", "orange", "yellow", "green", "blue", "indigo", "violet")

for color in rainbow:
    print(color)
```

We can see that the indented part of the `for` loop, known as the body, is executed once for each item in `rainbow`. The name `color` is assigned to the next value in `rainbow` at the start of each iteration. Note that the name `color` is arbitrary; we could easily have named it something else. The important thing is we stay consistent throughout the `for` loop. 

```python
for another_name in rainbow:
    print(another_name)
```

In general, however, we would like the variable name to be somewhat informative. 

**Question 6.** In the following cell, we've loaded the text of _Pride and Prejudice_ by Jane Austen, split it into individual words, and stored these words in an array `p_and_p_words`. Use a `for` loop to load each word in the array `p_and_p_words` into a variable one at a time, and check if the word is longer than 5 letters. When the loop completes, `longer_than_five` should contain an intger that represents the number of words in _Pride and Prejudice_ that are more than 5 letters long.

**Hint:** You can find the number of letters in a word/string with the `len` function. You can reference the code cell below for an example:

```python
my_word = 'hello'
len(my_word)
```

```python
austen_string = open('data/Austen_PrideAndPrejudice.txt', encoding='utf-8').read()
p_and_p_words = np.array(austen_string.split())
longer_than_five = ...

# a for loop would be useful here
longer_than_five
```

```python
grader.check("q6")
```

**Question 7.** Complete a simulation with 10,000 trials, by writing a for loop to pick two words uniformly at random (with replacement) from _Pride and Prejudice_ and determine if the words have different lengths. When the simulation is completed, `num_different` should be equal to the number of times that the two words selected in the simulations were different in length.

**Hints:** 

- What function did we use in section 1 to sample at random with replacement from an array? 

- Remember that `!=` checks for non-equality between two items.

```python
trials = 10000
num_different = ...

for ... in ...:
    ...
num_different
```

```python
grader.check("q7")
```

**Question 8.** Allie is playing darts. Her dartboard contains ten equal-sized zones with point values from 1 to 10. Write code that simulates her total score after 1000 dart tosses. You should **NOT** need a for loop to complete this problem. Remember that `np.random.choice(...)` can select multiple values from the same array if you provide a second argument that is an integer.

```python
possible_point_values = ...
num_tosses = 1000
simulated_tosses = ...
total_score = ...
total_score
```

```python
grader.check("q8")
```

## 3. Sampling Basketball Data

We will now introduce the topic of sampling, which we’ll be discussing in more depth in this week’s lectures. We’ll guide you through this code, but if you wish to read more about different kinds of samples before attempting this question, you can check out [Section 10 of the textbook](https://www.inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html).

Run the cell below to load player and salary data that we will use for our sampling. 

```python
player_data = Table().read_table("data/player_data.csv")
salary_data = Table().read_table("data/salary_data.csv")
full_data = salary_data.join("PlayerName", player_data, "Name")

# The show method immediately displays the contents of a table. 
# This way, we can display multiple tables using a single cell.
player_data.show(3)
salary_data.show(3)
full_data.show(3)
```

Rather than getting data on every player (as in the tables loaded above), imagine that we had gotten data on only a smaller subset of the players. For 492 players, it's not so unreasonable to expect to see all the data, but usually we aren't so lucky. 

If we want to make estimates about a certain numerical property of the population (known as a statistic, e.g. the mean or median), we may have to come up with these estimates based only on a smaller sample. Whether these estimates are useful or not often depends on how the sample was gathered. We have prepared some example sample datasets to see how they compare to the full NBA dataset. Later we'll ask you to create your own samples to see how they behave.

To save typing and increase the clarity of your code, we will package the analysis code into a few functions. This will be useful in the rest of the lab as we will repeatedly need to create histograms and collect summary statistics from that data.

We've defined the `histograms` function below, which takes a table with columns `Age` and `Salary` and draws a histogram for each one. It uses bin widths of 1 year for `Age` and $1,000,000 for `Salary`.

**Note:** This function ONLY takes in a table with these labels. You use this function providing a table as the input: for example `histograms(full_data)`

```python
def histograms(t):
    ages = t.column('Age')
    salaries = t.column('Salary')/1000000
    t1 = t.drop('Salary').with_column('Salary', salaries)
    age_bins = np.arange(min(ages), max(ages) + 2, 1) 
    salary_bins = np.arange(min(salaries), max(salaries) + 1, 1)
    t1.hist('Age', bins=age_bins, unit='year')
    plt.title('Age distribution')
    t1.hist('Salary', bins=salary_bins, unit='million dollars')
    plt.title('Salary distribution') 
    
histograms(full_data)
print('Two histograms should be displayed below')
```

**Question 9.** Create a function called `compute_statistics` that takes a table containing ages and salaries that does 3 things:

1. Draws a histogram of ages

2. Draws a histogram of salaries

3. Returns a two-element array containing the average age and average salary (in that order)

You can call the `histograms` function that was defined earlier to draw the histograms to make things a bit easier.

**Note:** More charts will be displayed when running the test cell. Please feel free to ignore the charts.

```python
def compute_statistics(age_and_salary_data):
    ...
    age = ...
    salary = ...
    ...

full_stats = compute_statistics(full_data)
full_stats
```

```python
grader.check("q9")
```

### Convenience sampling
One sampling methodology, which is **generally a bad idea**, is to choose players who are somehow convenient to sample.  For example, you might choose players from one team who are near your house, since it's easier to survey them.  This is called, somewhat pejoratively, *convenience sampling*.

Suppose you survey only *relatively new* players with ages less than 22.  (The more experienced players didn't bother to answer your surveys about their salaries.)

**Question 10.** Assign `convenience_sample` to a subset of `full_data` that contains only the rows for players under the age of 22.

```python
convenience_sample = ...
convenience_sample
```

```python
grader.check("q10")
```

**Question 11.** Assign `convenience_stats` to an array of the average age and average salary of your convenience sample, using the `compute_statistics` function.  Since they're computed on a sample, these are called *sample averages*. 

```python
convenience_stats = ...
convenience_stats
```

```python
grader.check("q11")
```

Next, we'll compare the convenience sample salaries with the full data salaries in a single histogram. To do that, we'll need to use the `bin_column` option of the `hist` method, which indicates that all columns are counts of the bins in a particular column. 

**The following cell does not require any changes; just run it**.

```python
def compare_salaries(first, second, first_title, second_title):
    """Compare the salaries in two tables."""
    first_salary_in_millions = first.column('Salary')/1000000
    second_salary_in_millions = second.column('Salary')/1000000
    first_tbl_millions = first.drop('Salary').with_column('Salary', first_salary_in_millions)
    second_tbl_millions = second.drop('Salary').with_column('Salary', second_salary_in_millions)
    max_salary = max(np.append(first_tbl_millions.column('Salary'), second_tbl_millions.column('Salary')))
    bins = np.arange(0, max_salary+1, 1)
    first_binned = first_tbl_millions.bin('Salary', bins=bins).relabeled(1, first_title)
    second_binned = second_tbl_millions.bin('Salary', bins=bins).relabeled(1, second_title)
    first_binned.join('bin', second_binned).hist(bin_column='bin', unit='million dollars')
    plt.title('Salaries for all players and convenience sample')

compare_salaries(full_data, convenience_sample, 'All Players', 'Convenience Sample')
```

<!-- BEGIN QUESTION -->

**Question 12.** Does the convenience sample give us an accurate picture of the salary of the full population? Would you expect it to, in general?  Why, why not? Before you move on, write a short answer in English below.  You can refer to the statistics calculated above or perform your own analysis.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Simple random sampling
A more justifiable approach is to sample uniformly at random from the players.  In a **simple random sample (SRS) without replacement**, we ensure that each player is selected at most once. Imagine writing down each player's name on a card, putting the cards in an box, and shuffling the box.  Then, pull out cards one by one and set them aside, stopping when the specified sample size is reached.

### Producing simple random samples
Sometimes, it’s useful to take random samples even when we have the data for the whole population. It helps us understand sampling accuracy.

### `sample`

The table method `sample` produces a random sample from the table. By default, it draws at random **with replacement** from the rows of a table. It takes in the sample size as its argument and returns a **table** with only the rows that were selected. 

Run the cell below to see an example call to `sample()` with a sample size of 5, with replacement.

```python
# Just run this cell
salary_data.sample(5)
```

The optional argument `with_replacement=False` can be passed through `sample()` to specify that the sample should be drawn without replacement.

Run the cell below to see an example call to `sample()` with a sample size of 5, without replacement.

```python
# Just run this cell
salary_data.sample(5, with_replacement = False)
```

<!-- BEGIN QUESTION -->

**Question 13.** Produce a simple random sample of size 44 from `full_data`. Run your analysis on it again.  Run the cell a few times to see how the histograms and statistics change across different samples.

- How much does the average age change across samples? 

- What about average salary?

_Type your answer here, replacing this text._

```python
my_small_srswor_data = ...
my_small_stats = ...
my_small_stats
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 14.**  Using the same method as the previous question, analyze several simple random samples but this time of size 100 from `full_data`.  

- Do the histogram shapes seem to change more or less across samples of 100 than across samples of size 44?  

- Are the sample averages and histograms closer to their true values/shape for age or for salary?  What did you expect to see?

_Type your answer here, replacing this text._

```python
# You can reuse your code from question 3.5 and tweak it
# use a sample size of 100 players

my_large_srswor_data = ...
my_large_stats = ...
my_large_stats
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the left side of the screen, right-clicking on it, and selecting **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 0.5)
```
>>> number_cheese == 3
True
```


Tests for q10 (points: None)
```
>>> convenience_sample.num_columns
11
```
```
>>> convenience_sample.num_rows
44
```


Tests for q11 (points: None)
```
>>> len(convenience_stats)
2
```
```
>>> round(float(convenience_stats[0]), 2) == 20.36
True
```
```
>>> round(float(convenience_stats[1]), 2) == 2383533.82
True
```


Tests for q2 (points: None)
```
>>> say_please == 'More please'
True
```


Tests for q3 (points: None)
```
>>> nacho_reaction('salsa')
'Spicy!'
```
```
>>> nacho_reaction('cheese')
'Cheesy!'
```
```
>>> nacho_reaction('both')
'Wow!'
```
```
>>> nacho_reaction('neither')
'Meh.'
```


Tests for q4 (points: None)
```
>>> np.count_nonzero(ten_nachos_reactions.column('Reactions') == make_array('Meh.', 'Cheesy!', 'Wow!', 'Wow!', 'Cheesy!', 'Spicy!', 'Wow!', 'Meh.', 'Cheesy!', 'Wow!')) == 10
True
```


Tests for q5 (points: None)
```
>>> 2 < number_wow_reactions < 6
True
```
```
>>> number_wow_reactions == 4
True
```


Tests for q6 (points: None)
```
>>> longer_than_five == 35453
True
```


Tests for q7 (points: None)
```
>>> 8100 <= num_different <= 9100
True
```


Tests for q8 (points: None)
```
>>> 1000 <= total_score <= 10000
True
```


Tests for q9 (points: None)
```
>>> stats = compute_statistics(full_data)
>>> plt.close()
>>> plt.close()
>>> round(float(stats[0]), 2) == 26.54
True
```
```
>>> stats = compute_statistics(full_data)
>>> plt.close()
>>> plt.close()
>>> round(float(stats[1]), 2) == 4269775.77
True
```
