```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

# Lab 07: Helicopter Drop

Welcome to Lab 07!

In this lab we will explore a scenario that is suitable for A/B testing. A/B testing is a form of hypothesis testing that allows you to make comparisons between two distributions.

For all problems that you must write our explanations and sentences for, you **must** provide your answer in the designated space. Also, throughout this lab please be sure to not re-assign variables throughout the notebook. For example, if you use `expected_proportion_correct` in your answer to one question, do not reassign it later on. Moreover, be mindful that functions that you write in earlier cells will need to be referenced to answer questions in later questions.


**Collaboration Policy:** Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

First, set up the tests and imports by running the cell below.

```python
# Run this cell, but please don't change it.
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 1. A/B Testing

You'll almost never be explicitly asked to perform an A/B test. Make sure you can identify situations where the test is appropriate and know how to correctly implement each step.

**Question 1.** The following statements are the unordered steps of an A/B hypothesis test:

1. Choose a test statistic (typically the difference in means between two categories)

2. Shuffle the labels of the original sample, find your simulated test statistic, and repeat many times

3. Find the value of the observed test statistic

4. Calculate the p-value based off your observed and simulated test statistics

5. Define a null and alternate model

6. Use the p-value and p-value cutoff to draw a conclusion about the null hypothesis

Make an array called `ab_test_order` that contains the correct order of an A/B test, where the first item of the array is the first step of an A/B test and the last item of the array is the last step of an A/B test


```python
ab_test_order = ...
```

```python
grader.check("q1")
```

<!-- BEGIN QUESTION -->

**Question 2.** If the null hypothesis of an A/B test is correct, should the order of labels affect the differences in means between each group?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 3.** Why do we shuffle labels in an A/B test? 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Paper Helicopters

The Paper Helicopter Experiment (https://www.paperhelicopterexperiment.com/) provides different templates for constructing paper helicopters that can be dropped. The designs can be modified to create helicopters with different dimensions for the wings, body, weighting, and more. In this lab you'll investigate whether the length of the rotor and the amount of weight attached to helicopter results in a different in the time it takes to fall from a 3rd floor stairwell. Do conduct your analysis, you'll use an A/B test.

### The Data

The data was collected by a group of NCSSM Online students during online weekend held October 2nd, 2021. Five groups of students dropped 143 paper helicopters in total. There were 4 unique configurations that were assigned to the groups, one per group.

1. **The Default:** Full length (long) rotors, unweighted body

2. **The Short:** Short rotors, unweighted body

3. **The Heavy:** Long rotors, weighted body

4. **The Short Heavy:** Short rotors, weighted body

The following features were recorded for each helicopter dropped:
* `Rotor`: Long or Short (string). Indicates if the helicopter had full length or shortened rotors.
* `Weight`: Weighted or Unweighted (string). Indicates if the helicopter had a weighted (3 paperclips) or unweighted body.
* `Team`: 1, 2, 3, 4, or 5 (int). Indicates which team
* `Location`: Side A or Side B (string). Indicates which side of the building the helicopters were dropped from.
* `Obstruction`: True or False (Boolean). Indicates if the helicopter hit the stairwell or other object on the way down.
* `Time`: (Float). Time, measured in seconds, it took for the helicopter to hit the ground after being released.

Run the following cell to load the data in as the table named `copters`.

```python
copters = Table.read_table('data/helicopters.csv')
copters
```

Due to variation in the construction, weather, and randomness of the universe not every drop for each configuration was identical. Let's explore the data a bit to see how the drop times varied.

**Question 4.** Let's start by cleaning the data a little. This dataset has several measurements taken where the helicopter hit the stairwell, a nearby building, or some other obstruction that impacted the fall time. These results should be discarded since they wouldn't be an accurate reflection of what we're trying to measure.

Create a new table called `helicopters` that only contains the rows of `copters` that correspond to a helicopter that did not hit an obstruction during the fall.


```python
helicopters = ...
helicopters
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

**Question 5.** Suppose you're interested in determining how the rotor length might impact fall time. Start by creating a histogram that displays the distribution of times for helicopters with long rotors. Use the provided bins stored to `my_bins`.


```python
my_bins = np.arange(2, 9.5, 0.5)
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 6.** Now, construct a histogram that shows the distribution of times for only helicopters with short rotors. Again, use the provided bins.


```python
my_bins = np.arange(2, 9.5, 0.5)
...
```

<!-- END QUESTION -->

It can be difficult sometimes to compare two histograms when they are on separate graphs. However, by providing the argument `group` to the `.hist()` method with the label (string) for a column you wish to group by, and it will produce a histogram for each unique group.

For example, `table.hist('col_label', bins=your_bins, group='feature')` would produce a histogram for each unique group found in the `feature` column of `table`, using the numerical data in the column labeled `col` label.

Running the cell below will create 2 histograms, one for each group as determined by the values in the `Rotor` column.

```python
my_bins = np.arange(2, 9.5, 0.5)
helicopters.hist('Time', bins = my_bins, group = "Rotor" )
```

You can see it's much easier to compare these two distributions now that they're on the same scale and set of axes.

## 3. The Rotor Question

You might now be wondering if there is a difference between the fall times of helicopters with long (full length) vs. short rotors. It certainly appears that the distributions might be different than each other. We'll use our data to attempt to draw a conclusion.

<!-- BEGIN QUESTION -->

**Question 7.** We want to know whether shortening the rotor length **causes** a change in the fall time.  

Why is it not sufficient to compare fall times of the helicopters with short rotors and those with long rotors?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8.** Why is an A/B test appropriate for this situation? What would you use as the "A" group (the control group) and which group is the "B" group (the treatment group)?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 9.**  Create a table `time_means` that contains the average fall time for the full length rotor helicopters and the shortened rotor helicopters. It should have two columns: one indicating the rotor length, and one that contains the average fall time for each group.

**Hint:** Use a combination of `.select()` and `.group()` to make short work of this.


```python
time_means = ...
time_means
```

```python
grader.check("q9")
```

<!-- BEGIN QUESTION -->

**Question 10.** We want to figure out if there is a difference between the distribution of fall times for full length and shortened rotor helicopters. Specifically, we want to test if fall times were longer for longer rotor helicopters than they were for shorter rotor helicopters. 

What should the test statistic be? How does it help us differentiate whether the data supports the null and alternative? 

Confirm your answer with your instructor (or a classmate who already has) before moving on.

**Note:** General statements like "the test statistic will help us determine if there is a difference" or " This will help us differentiate where the data supports the null and alternative", will be considered incomplete responses.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 11.** Calculate the test statistic for our observed data using the `time_means` table. Set `observed_difference` to the value you compute. 


```python
observed_difference = ...
observed_difference
```

```python
grader.check("q11")
```

**Question 12.** To perform an A/B test, we'll need to compute similar statistics for our many simulations we'll run. Writing a function which returns the statistic will be a great way to save time. Write a function named `find_test_stat` which takes in the arguments `table`, `labels_col`, and `values_col` that calculates the test statistic for this problem. The `table` passed into this function will be a permutation of our original table and structured the same way. `labels_col` will be passed as a string that specifies which column in the table contains the label you'll be grouping by. `values_col` specifies which column contains the values that you'll be using to compute the test statistic.

When you've written this function, you can pass it any table and the two specified labels and it will compute a test statistic for an A/B test, not just for this problem, but any problem! A very handy function to keep around. 

Running `find_test_stat(helicopters, "Rotor", "Time")` should return the exact same value your code from the previous question computed: the difference between of average values in the "Time column between the groups determined by "Rotor" column in the `helicopters` table. However, you should also be able to run `find_test_stat(helicopters, "Weight", "Time")` to compute the difference between of average values in the **"Weight"** column between the groups determined by "Rotor" column in the `helicopters` table just as easily. 


```python
def find_test_stat(table, labels_col, values_col):
    ...
find_test_stat(helicopters, "Rotor", "Time")
```

```python
grader.check("q12")
```

When we run a simulation for A/B testing, we resample the data by shuffling the labels of the original sample. If the null hypothesis is true and the time distributions **are** the same, we expect that the difference in mean time will be not change when "Rotor" labels are changed.

**Question 13.** Write a function `simulate_and_test_statistic` to compute one trial of our A/B test. Your function should run a simulation and return a test statistic.

**Hint:** You can "shuffle" the labels by using `.sample(with_replacement = False)` on the table, and then just select the column (as an array using `.column`) that contains the newly shuffled labels. After that, you can either overwrite the existing labels, or, extend the table with a new column labeled something like "shuffled labels". Just make sure you pass the correct label on to `find_test_stat`!

**Note:** The autograder test here is fairly lenient, if you have an issue with the following questions, make sure to take a look at your answer to the previous question. Specifically, make sure that you are taking the directionality of our alternative hypothesis into account, meaning, check the order in which you are subtracting the average times.


```python
def simulate_and_test_statistic(table, labels_col, values_col):
    ...
    
simulate_and_test_statistic(helicopters, "Rotor", "Time")
```

```python
grader.check("q13")
```

**Question 14.** Now that we can quickly shuffle the labels and compute a test statistic, simulate 5000 trials of our A/B test and store the test statistics in an array called `differences`.


```python
# This cell might take a couple seconds to run
differences = make_array()
...
differences
```

```python
grader.check("q14")
```

Run the cell below to view a histogram of your simulated test statistics plotted with your observed test statistic.

```python
Table().with_column('Difference Between Group Means', differences).hist()
plt.scatter(observed_difference, -0.01, color = 'red', s = 60, zorder = 2);
```

**Question 15.** Find the p-value for your test and assign it to `empirical_P`.


```python
empirical_P = ...
empirical_P
```

```python
grader.check("q15")
```

<!-- BEGIN QUESTION -->

**Question 16.** Using a 5% P-value cutoff, draw a conclusion about the null and alternative hypotheses. Describe your findings using simple, non-technical language. What does your analysis tell you about fall times of helicopters with full length vs shortened rotors? What can you claim about causation from your statistical analysis?


_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> len(ab_test_order)
6
```
```
>>> isinstance(ab_test_order, np.ndarray)
True
```


Tests for q11 (points: None)
```
>>> isinstance(observed_difference, float)
True
```
```
>>> round(observed_difference, 3)
1.647
```


Tests for q12 (points: None)
```
>>> round(find_test_stat(helicopters, 'Rotor', 'Time'), 3)
1.647
```
```
>>> round(find_test_stat(helicopters, 'Weight', 'Time'), 3)
1.0
```


Tests for q13 (points: None)
```
>>> np.random.seed(123)
>>> 0 < round(simulate_and_test_statistic(helicopters, 'Rotor', 'Time'), 3) < 1
True
```


Tests for q14 (points: None)
```
>>> len(differences)
5000
```
```
>>> abs(np.average(differences)) < 1
True
```
```
>>> all(differences == differences.item(0)) == False
True
```


Tests for q15 (points: None)
```
>>> empirical_P < 0.05
True
```


Tests for q4 (points: None)
```
>>> helicopters.num_rows
130
```
```
>>> all(helicopters.column('Obstruction') == False)
True
```


Tests for q9 (points: None)
```
>>> time_means.num_rows
2
```
```
>>> time_means.where('Rotor', 'Long').column(1).item(0)
5.610666666666667
```
```
>>> time_means.where('Rotor', 'Short').column(1).item(0)
3.963636363636364
```
