```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 01: Causality and Expressions

Please complete this notebook by filling in the cells provided. Before you begin, be sure you've run the very first cell in this notebook to load `otter`, the auto-grading library that will provide feedback while completing this assignment. This cell should be run each time you load this notebook, as everything stored into memory is reset when you close Datahub for more than a few minutes.

**Recommended Reading:**

- [What is Data Science](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!


**Notes:** 

- Start early so that you can come to office hours if you're stuck. Check the course website for the office hours schedule. Late work will not be accepted as per the course expectations. Late work will not be accepted without advanced permission as per the course expectations document, so plan accordingly. You should start early so that you have time to get help if you're stuck. Office hours are held throughout the week, so be sure to use them if needed.

- This homework has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. Additional tests will be run once your homework is submitted to Gradescope**. Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and not super far off the expected answer. The tests that are run after submission will evaluate accuracy more carefully. Do not assume that just because all your tests pass means that your answers are correct!

- Directly sharing answers is **not** okay, but discussing problems with the course staff (teacher and TAs) or with other students is encouraged. Refer to the course expectations document to learn more about how to learn and work cooperatively.

## 1. Scary Arithmetic

An ad for ADT Security Systems says,

> "When you go on vacation, burglars go to work [...] According to FBI statistics, over 25% of home burglaries occur between Memorial Day to Labor Day."

<!-- BEGIN QUESTION -->

**Question 1.** This ad is suggesting that you should be worried about going on your summer vacation because there's an increased chance that your home will be broken into during this time period. Do the data in the ad support the claim that burglars are more likely to go to work during the time between Memorial Day to Labor Day than other times of the year? Put another way, should you believe the security company that you should be on alert during the summer for an increased risk of burglary? Please explain your answer.

**Note:** You can assume that "over 25%" means only slightly over. Had it been much over, say closer to 30%, then the marketers would have said so.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 2. Characters in Little Women

We can use some Python code to count the number of times that the literary characters were named in each chapter of the classic book, [*Little Women*](https://inferentialthinking.com/chapters/01/3/1/Literary_Characters.html). In computer science, the word "character" also refers to a letter, digit, space, or punctuation mark; any single element of a text. The following code generates a scatter plot in which each dot corresponds to a chapter of *Little Women*. The horizontal position of a dot measures the number of periods in the chapter. The vertical position measures the total number of characters.

```python
# This cell contains code that hasn't yet been covered in the course,
# but you should be able to interpret the scatter plot it generates.

from datascience import *
from urllib.request import urlopen
import numpy as np
%matplotlib inline

little_women_url = 'https://www.inferentialthinking.com/data/little_women.txt'
chapters = urlopen(little_women_url).read().decode().split('CHAPTER ')[1:]
text = Table().with_column('Chapters', chapters)
Table().with_columns(
    'Periods',    np.char.count(chapters, '.'),
    'Characters', text.apply(len, 0)
    ).scatter(0)
```

**Question 2.** Around how many periods are there in the chapter with the most characters? Assign either 1, 2, 3, 4, or 5 to the name `characters_q2` below.

1. 250

1. 390

1. 440

1. 32,000

1. 40,000


```python
characters_q2 = ...
```

```python
grader.check("q2")
```

The test above checks that your answers are in the correct format and that you have the correct answer. **The questions that follow do not contain tests to check that you answered correctly**, only that you assigned a number successfully in each multiple-choice answer cell.

**Question 3.** Which of the following chapters has the most characters per period? Assign either 1, 2, or 3 to the name `characters_q3` below.

1. The chapter with about 60 periods

1. The chapter with about 350 periods

1. The chapter with about 440 periods


```python
characters_q3 = ...
```

```python
grader.check("q3")
```

**Again**, the test above checks that your answers are in the correct format, but not that you have answered correctly.

To discover more interesting facts from this plot, read [Section 1.3.2](https://inferentialthinking.com/chapters/01/3/2/Another_Kind_Of_Character.html) of the textbook.

## 3. Names and Assignment Statements

When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

**Question 4.**  Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q4` below to indicate your answer.

1. Python is smart and already knows `4 = 2 + 2`.

1. `4` is already a defined number, and it doesn't make sense to make a number be a name for something else. In Python, "`x = 2 + 2`" means "assign `x` as the name for the value of `2 + 2`."

1. It should be `2 + 2 = 4`.

1. I don't get an error message. This is a trick question.


```python
names_q4 = ...
```

```python
grader.check("q4")
```

When you run the following cell, Python will produce another cryptic error message.

```python
two = 3
six = two plus two
```

**Question 5.**  Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q5` below to indicate your answer.

1. The `plus` operation only applies to numbers, not the word "two".

1. The name "two" cannot be assigned to the number 3.

1. Two plus two is four, not six.

1. Python cannot interpret the name `two` followed directly by a name that has not been defined.


```python
names_q5 = ...
```

```python
grader.check("q5")
```

When you run the following cell, Python will, yet again, produce another cryptic error message.

```python
x = print(5)
y = x + 2
```

**Question 6.** Choose the best explanation of what's wrong with the code and assign 1, 2, or 3 to `names_q6` below to indicate your answer.

1. Python doesn't want `y` to be assigned.

1. The `print` operation is meant for displaying values to the programmer, not for assigning values!

1. Python can’t do addition between one name and one number. It has to be 2 numbers or 2 predefined names.


```python
names_q6 = ...
```

```python
grader.check("q6")
```

## 4. Job Opportunities & Education in Rural India

A [study](http://www.nber.org/papers/w16021.pdf) at UCLA investigated factors that might result in greater attention to the health and education of girls in rural India. One such factor is information about job opportunities for women. The idea is that if people know that educated women can get good jobs, they might take more care of the health and education of girls in their families, as an investment in the girls’ future potential as earners. Without the knowledge of job opportunities, the author hypothesizes that families do not invest in women’s well-being.

The study focused on 160 villages outside the capital of India, all with little access to information about call centers and similar organizations that offer job opportunities to women. In 80 of the villages chosen at random, recruiters visited the village, described the opportunities, recruited women who had some English language proficiency and experience with computers, and provided ongoing support free of charge for three years. In the other 80 villages, no recruiters visited and no other intervention was made.

At the end of the study period, the researchers recorded data about the school attendance and health of the children in the villages.

**Question 7.** Which statement best describes the *treatment* and *control* groups for this study? Assign either 1, 2, or 3 to the name `jobs_q7` below.

1. The treatment group was the 80 villages visited by recruiters, and the control group was the other 80 villages with no intervention.

1. The treatment group was the 160 villages selected, and the control group was the rest of the villages outside the capital of India.

1. There is no clear notion of *treatment* and *control* group in this study.


```python
jobs_q7 = ...
```

```python
grader.check("q7")
```

**Question 8.** Was this an observational study or a randomized controlled experiment? Assign either 1, 2, or 3 to the name `jobs_q8` below.

1. This was an observational study.

1. This was a randomized controlled experiment.  

1. This was a randomized observational study.


```python
jobs_q8 = ...
```

```python
grader.check("q8")
```

<!-- BEGIN QUESTION -->

**Question 9.** The study reported, "Girls aged 5-15 in villages that received the recruiting services were 3 to 5 percentage points more likely to be in school and experienced an increase in Body Mass Index, reflecting greater nutrition and/or medical care. However, there was no net gain in height. For boys, there was no change in any of these measures." Why do you think the author points out the lack of change in the boys? Please explain your answer.

**Hint:** Remember the original hypothesis. The author believes that educating women in job opportunities will cause families to invest more in the women's well-being.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 5. Differences in Math Placement


NCSSM’s Math Department keeps records regarding incoming junior's initial Math course. The table below displays the numbers of first semester juniors in each of the listed courses academic years 2010-2011 and 2020-2021.

| Course                             | 2010-2011    | 2020-2021   |
|------------------------------------|--------------|-------------|
| Precalculus                        |      162     |    63       |
| Calculus                           |      72      |    175      |
| Multivariable Calculus             |      18      |    53       |



**Question 10.** Suppose you want to find the **biggest absolute difference** between the numbers of students in the two years, among the three courses.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 18 from 53) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers


```python
biggest_change = ...
biggest_change
```

```python
grader.check("q10")
```

**Question 11.** Which of the three courses had the **smallest absolute difference**? Assign `smallest_change_course` to 1, 2, or 3 where each number corresponds to the following course:

1. Precalculus  

1. Calculus  

1. Multivariable Calculus

Choose the number that corresponds to the course with the smallest absolute difference.

You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each course. 


```python
smallest_change_course = ...
smallest_change_course
```

```python
grader.check("q11")
```

**Question 12.** For each course, define the relative change to be the following:

$$\frac{\text{absolute difference}}{\text{value in 2010-2011}} \cdot 100$$ 

Fill in the code below such that `precalculus_relative_change`, `calculus_relative_change` and `multi_relative_change` are assigned to the relative changes for their respective courses.


```python
precalculus_relative_change = (abs(...) / 162) * 100
calculus_relative_change = ...
multi_relative_change = ...
precalculus_relative_change, calculus_relative_change, multi_relative_change
```

```python
grader.check("q12")
```

**Question 13.** Assign `biggest_rel_change_course` to 1, 2, or 3 where each number corresponds to to the following: 

1. Precalculus 

1. Calculus  

1. Multivariable Calculus

Choose the number that corresponds to the course with the biggest relative change.


```python
# Assign biggest_rel_change_major to the number corresponding to the major with the biggest relative change.
biggest_rel_change_course = ...
biggest_rel_change_course
```

```python
grader.check("q13")
```

## 6. Nearsightedness Study


Myopia, or nearsightedness, results from a number of genetic and environmental factors. In 1999, Quinn et al studied the relation between myopia and ambient lighting at night (for example, from nightlights or room lights) during childhood.

<!-- BEGIN QUESTION -->

**Question 14.** The data were gathered by the following procedure, reported in the study. 

> "Between January and June 1998, parents of children aged 2-16 years [...] that were seen as outpatients in a university pediatric ophthalmology clinic completed a questionnaire on the child's light exposure both at present and before the age of 2 years." 

Was this study observational, or was it a controlled experiment? Explain. 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 15.** The study found that of the children who slept with a room light on before the age of 2, 55% were myopic. Of the children who slept with a night light on before the age of 2, 34% were myopic. Of the children who slept in the dark before the age of 2, 10% were myopic. The study concluded that, "The prevalence of myopia [...] during childhood was strongly associated with ambient light exposure during sleep at night in the first two years after birth."

Do the data support this statement? You may interpret "strongly" in any reasonable qualitative way.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 16.** On May 13, 1999, CNN reported the results of this study under the headline, "Night light may lead to nearsightedness." Does the conclusion of the study claim that the use of a night light causes nearsightedness?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 17.** The final paragraph of the CNN report said that "several eye specialists" had pointed out that the study should have accounted for heredity.

Myopia is passed down from parents to children. Myopic parents are more likely to have myopic children, and may also be more likely to leave lights on habitually (since the parents have poor vision). In what way does the knowledge of this possible genetic link affect how we interpret the data from the study? Explain.


_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q10 (points: None)
```
>>> isinstance(biggest_change, (int, float, float))
True
```


Tests for q11 (points: None)
```
>>> isinstance(smallest_change_course, int)
True
```


Tests for q12 (points: None)
```
>>> isinstance(precalculus_relative_change, float)
True
```
```
>>> isinstance(calculus_relative_change, float)
True
```
```
>>> isinstance(multi_relative_change, float)
True
```
```
>>> multi_relative_change > calculus_relative_change > precalculus_relative_change > 50
True
```


Tests for q13 (points: None)
```
>>> biggest_rel_change_course in [1, 2, 3]
True
```


Tests for q2 (points: None)
```
>>> characters_q2 != ...
True
```
```
>>> characters_q2 in [1, 2, 3, 4, 5]
True
```
```
>>> characters_q2 != 3
True
```
```
>>> characters_q2 == 2
True
```


Tests for q3 (points: None)
```
>>> characters_q3 != ...
True
```
```
>>> characters_q3 in [1, 2, 3]
True
```


Tests for q4 (points: None)
```
>>> names_q4 in [1, 2, 3, 4]
True
```


Tests for q5 (points: None)
```
>>> names_q5 in [1, 2, 3, 4]
True
```


Tests for q6 (points: None)
```
>>> names_q6 in [1, 2, 3]
True
```


Tests for q7 (points: None)
```
>>> jobs_q7 in [1, 2, 3]
True
```


Tests for q8 (points: None)
```
>>> jobs_q8 in [1, 2, 3]
True
```
