```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Run the cell below

To run a code cell (i.e.; execute the python code inside a Jupyter notebook) you can click the play button on the ribbon underneath the name of the notebook. Before you begin click the "Run cell" button at the top that looks like ▶| or hold down `Shift` + `Return`.

# Homework 05: Probability, Simulation, Estimation, and Assessing Models

```python
# Don't change this cell; just run it. 
import numpy as np
from datascience import *

import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

**Recommended Reading**:

* [Randomness](https://inferentialthinking.com/chapters/09/Randomness.html) 

* [Sampling and Empirical Distributions](https://inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)

* [Testing Hypotheses](https://inferentialthinking.com/chapters/11/Testing_Hypotheses.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Notes:** 

Late work will not be accepted as per the course expectations. Late work will not be accepted without advanced permission as per the course expectations document, so plan accordingly. You should start early so that you have time to get help if you're stuck. Office hours are held throughout the week, so be sure to use them if needed.

- This homework has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. Additional tests will be run once your homework is submitted to Gradescope**. Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and not super far off the expected answer. The tests that are run after submission will evaluate accuracy more carefully. Do not assume that just because all your tests pass means that your answers are correct!

- Directly sharing answers is **not** okay, but discussing problems with the course staff (teacher and TAs) or with other students is encouraged. Refer to the course expectations document to learn more about how to learn and work cooperatively.

## 1. Probability


We will be testing some probability concepts that were introduced in class. For all of the following problems, we will introduce a problem statement and give you a proposed answer. You must assign the provided variable to one of the following three integers, depending on whether the proposed answer is too low, too high, or correct. 

1. Assign the variable to 1 if you believe our proposed answer is too high.

1. Assign the variable to 2 if you believe our proposed answer is too low.

1. Assign the variable to 3 if you believe our proposed answer is correct.


You are more than welcome to create more cells across this notebook to use for arithmetic operations.

**Question 1.** You roll a 6-sided die 10 times. What is the chance of getting 10 sixes?

Our proposed answer: $$\left(\frac{1}{6}\right)^{10}$$

Assign `ten_sixes` to either 1, 2, or 3 depending on if you think our answer is too high, too low, or correct. 

```python
ten_sixes = ...
ten_sixes
```

```python
grader.check("q1")
```

**Question 2.** Take the same problem set-up as before, rolling a fair dice 10 times. What is the chance that every roll is less than or equal to 5?

Our proposed answer: $$1 - \left(\frac{1}{6}\right)^{10}$$

Assign `five_or_less` to either 1, 2, or 3. 


```python
five_or_less = ...
five_or_less
```

```python
grader.check("q2")
```

**Question 3.** Assume we are picking a lottery ticket. We must choose three distinct numbers from 1 to 1000 and write them on a ticket. Next, someone picks three numbers one by one from a bowl with numbers from 1 to 1000 each time without putting the previous number back in. We win if our numbers are all called in order. 

If we decide to play the game and pick our numbers as 12, 140, and 890, what is the chance that we win? 

Our proposed answer: $$\left(\frac{3}{1000}\right)^3$$

Assign `lottery` to either 1, 2, or 3. 


```python
lottery = ...
```

```python
grader.check("q3")
```

**Question 4.** Assume we have two lists, list A and list B. List A contains the numbers [20,10,30], while list B contains the numbers [10,30,20,40,30]. We choose one number from list A randomly and one number from list B randomly. What is the chance that the number we drew from list A is larger than or equal to the number we drew from list B?

Our proposed solution: $$1/5$$

Assign `list_chances` to either 1, 2, or 3. 

**Hint:** Consider the different possible ways that the items in List A can be greater than or equal to items in List B. Try working out your thoughts with a pencil and paper, what do you think the correct solutions will be close to?


```python
list_chances = ...
```

```python
grader.check("q4")
```

## 2. Monkeys Typing Shakespeare

A monkey is banging repeatedly on the keys of a typewriter. Each time, the monkey is equally likely to hit any of the 26 lowercase letters of the English alphabet, 26 uppercase letters of the English alphabet, and any number between 0-9 (inclusive), regardless of what it has hit before. There are no other keys on the keyboard.  

This question is inspired by a mathematical theorem called the Infinite [monkey theorem](https://en.wikipedia.org/wiki/Infinite_monkey_theorem), which postulates that if you put a monkey in the situation described above for an infinite time, they will eventually type out all of Shakespeare’s works.

**Question 5.** Suppose the monkey hits the keyboard 6 times.  Compute the probability that the monkey types the sequence of characters `ma4110`.  (Call this `ma4110_chance`.) Determine the theoretical probability (without simulation) and assign it as either fraction or arithmetic expression. For example your assignment statement could be formatted like `ma4110_chance = 1/10000` or `ma4110_chance = (1/100) ** 2` and still be graded correctly.


```python
ma4110_chance = ...
ma4110_chance
```

```python
grader.check("q5")
```

**Question 6.** Write a function called `simulate_key_strike`.  It should take **no arguments**, and it should return a random one-character string that is equally likely to be any of the 26 lower-case English letters, 26 upper-case English letters, or any number between 0-9 (inclusive). The provided code below will create a list called `keys` that contains all the lower-case English letters, upper-case English letters, and the digits 0-9 (inclusive).


```python
# Proivded code, do not change
import string
keys = list(string.ascii_lowercase + string.ascii_uppercase + string.digits)

def simulate_key_strike():
    """Simulates one random key strike."""
    # Your code goes below this line
    ...

# An example call to your function:
simulate_key_strike()
```

```python
grader.check("q6")
```

**Question 7.** Write a function called `simulate_several_key_strikes`.  It should take one argument: an integer specifying the number of key strikes to simulate. It should return a string containing that many characters, each one obtained from simulating a key strike by the monkey.

**Hint:** If you make a list or array of the simulated key strikes called `key_strikes_array`, you can convert that to a string by calling `"".join(key_strikes_array)`


```python
def simulate_several_key_strikes(num_strikes):
    ...

# An example call to your function:
simulate_several_key_strikes(11)
```

```python
grader.check("q7")
```

**Question 8.** Call `simulate_several_key_strikes` 5000 times, each time simulating the monkey striking 6 keys.  Compute the proportion of times the monkey types `"ma4110"`, calling that proportion `ma4110_proportion`.


```python
...
ma4110_proportion
```

```python
grader.check("q8")
```

<!-- BEGIN QUESTION -->

**Question 9.** Check the value your simulation computed for `ma4110_proportion`.  Is your simulation a good way to estimate the chance that the monkey types `"ma4110"` in 6 strikes (the answer to **Question 5**)?  Why or why not? Think about the theoretical probability of this event occurring and the number of simulations you ran when writing your response.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 10.** Compute the chance that the monkey types the letter `"m"` at least once in the 6 strikes.  Call it `m_chance`. Provide your answer as an expression Python can evalute to calculate the final probability. (For example, you should put `(1/6)**3` instead of `0.00462962962963`) 


```python
m_chance = ...
m_chance
```

```python
grader.check("q10")
```

<!-- BEGIN QUESTION -->

**Question 11.** Do you think that a computer simulation with 5000 trials would be more or less effective to estimate `m_chance` when compared to when we tried to estimate `ma4110_chance` this way? Why or why not? What specific criteria did you consider when making your decision?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 3. Sampling Basketball Players


This exercise uses salary data and game statistics for basketball players from the 2019-2020 NBA season. The data was collected from [Basketball-Reference](http://www.basketball-reference.com).

Run the next cell to load the two datasets.

```python
player_data = Table.read_table('data/player_data.csv')
salary_data = Table.read_table('data/salary_data.csv')
player_data.show(3)
salary_data.show(3)
```

**Question 12.** We would like to relate players' game statistics to their salaries.  Compute a table called `full_data` that includes one row for each player who is listed in both `player_data` and `salary_data`.  It should include all the columns from `player_data` and `salary_data`, except the `"Name"` column.

**Hint:** A `.join` operation would be helpful here to combine your tables.


```python
full_data = ...
full_data
```

```python
grader.check("q12")
```

Basketball team managers would like to hire players who perform well but don't command high salaries.  From this perspective, a very crude measure of a player's *value* to their team is the number of 3 pointers and free throws the player scored in a season for every **\$100000 of salary** 

**Note:** The `Salary` column is in dollars, not hundreds of thousands of dollars. 

For example, Al Horford scored an average of 5.2 points for 3 pointers and free throws combined, and has a salary of **\$28 million.** This is equivalent to 280 thousands of dollars, so his value is $\frac{5.2}{280}$. 

The formula used to make this calculation is:

$$\frac{\text{"PTS"} - 2 \times \text{"2P"}}{\text{"Salary"}\ / \ 100000}$$

<!-- BEGIN QUESTION -->

**Question 13.** Create a table called `full_data_with_value` that's a copy of `full_data`, with an extra column called `"Value"` containing each player's value (according to our crude measure).  Then make a histogram of players' values.  Use the specified bins, as they've been chose to make the histogram informative. Then, don't forget to include your units in the histogram! Remember that `hist()` takes in an optional third argument, `unit`, that allows you to specify the units of your data. Refer to the python reference sheet to look at `tbl.hist(...)` if necessary.

**Just so you know:** Informative histograms contain a majority of the data and **exclude outliers**. The provided bins will intentionally exclude some data points that are considered outliers.


```python
my_bins = np.arange(0, 0.7, .1) # Use these provided bins when you make your histogram
full_data_with_value = ...
...
```

<!-- END QUESTION -->

Now suppose we **weren't** able to find out every player's salary (perhaps it was too costly to interview each player).  Instead, we have gathered a *simple random sample* of 50 players' salaries.  The cell below will load a pre-made sample of 50 players to the table `sample_salary_data`.

```python
sample_salary_data = Table.read_table("data/sample_salary_data.csv")
sample_salary_data.show(3)
```

<!-- BEGIN QUESTION -->

**Question 14.** Make a histogram of the values of the players in `sample_salary_data`, using the same method for measuring value we used in question 3.2. Make sure to specify the same bins and units that were used in **Question 13**.

**Hint:** This will take several steps and perhaps several intermediate tables. Don't feel like you need to do this in a single line of code.


```python
sample_data = player_data.join('Player', sample_salary_data, 'Name')
sample_data_with_value = ...
...
```

<!-- END QUESTION -->



```python
sample_data
```

Now let us summarize what we have seen.  To guide you, we have written most of the summary already.

**Question 15.** Complete the statements below by setting each relevant variable name to the value that correctly fills the blank.

* The plot in **Question 13** displayed a(n) [`distribution_1`] distribution of the population of [`player_count_1`] players.  The areas of the bars in the plot sum to [`area_total_1`].

* The plot in **Question 14** displayed a(n) [`distribution_2`] distribution of the sample of [`player_count_2`] players.  The areas of the bars in the plot sum to [`area_total_2`].

`distribution_1` and `distribution_2` should be set to one of the following strings: `"empirical"` or `"probability"`. 

`player_count_1`, `area_total_1`, `player_count_2`, and `area_total_2` should be set to integers.

Remember that areas are represented in terms of **percentages**, not proportions.

**Hints:** 

- For a refresher on distribution types, check out [Section 10](https://inferentialthinking.com/chapters/10/Sampling_and_Empirical_Distributions.html)

- The `hist()` table method ignores data points outside the range of its bins, but you may ignore this fact and calculate the areas of the bars using what you know about histograms from class.


```python
distribution_1 = ...
player_count_1 = ...
area_total_1 = ...

distribution_2 = ...
player_count_2 = ...
area_total_2 = ...
```

```python
grader.check("q15")
```

<!-- BEGIN QUESTION -->

**Question 16.** For which range of values does the plot of the simple random sample in **Question 14** better represent the distribution of the **population's** player values: 0 to 0.3, or above 0.3? Explain your answer. 


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 4. Earthquakes


The next cell loads a table containing information about **every earthquake with a magnitude above 5** in 2019 (smaller earthquakes are generally not felt, only recorded by very sensitive equipment), compiled by the US Geological Survey). 

**Source:** https://earthquake.usgs.gov/earthquakes/search/

```python
earthquakes = Table().read_table('data/earthquakes_2019.csv').select(['time', 'mag', 'place'])
earthquakes
```

If we were studying all human-detectable 2019 earthquakes and had access to the above data, we’d be in good shape - however, if the USGS didn’t publish the full data, we could still learn something about earthquakes from just a smaller subsample. If we gathered our sample correctly, we could use that subsample to get an idea about the distribution of magnitudes (above 5, of course) throughout the year! 

We'll create two different samples from the `earthquake` table. Analyze the code for each sample to examine how the methods are similar and how they differ. You'll be asked to compare the results of each method, and understanding the differences in the sampling method will potentially help you explain any differences in the average value of `mag` that is computed.

### Sample #1

```python
# First sample method
sample1 = earthquakes.sort('mag', descending = True).take(np.arange(100))

# Calculate the mean of the first sample
sample1_magnitude_mean = np.mean(sample1.column('mag'))

sample1_magnitude_mean
```

### Sample #2

```python
# Second sample method
sample2 = earthquakes.take(np.arange(100))

# Calculate the mean of the second sample
sample2_magnitude_mean = np.mean(sample2.column('mag'))

sample2_magnitude_mean
```

<!-- BEGIN QUESTION -->

**Question 17.** Neither of these samples accurately represent the population from which they were drawn. Explain why each of the two samples would create a biased average of `mag`.

**Hint:** Consider the ordering of the rows in the tables from which the samples were constructed.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 18.** Write code to produce a sample of size 200 that **is** representative of the population and assign it to a Table named `representative_sample`. Then, take the mean of the magnitudes of the earthquakes in this sample and assign it to `representative_mean`. 

**Hint:** In class, you've learned what type of sample should be used to best represent a population.


```python
representative_sample = ...
representative_mean = ...
representative_mean
```

```python
grader.check("q18")
```

**Question 19.** Suppose we want to figure out what the biggest magnitude earthquake was in 2019, but we only have our representative sample of 200. Let’s see if trying to find the biggest magnitude in the population from a random sample of 200 is a reasonable idea!

In the cell below write code that uses simulation to create 5,000 random samples of size 200 from the `earthquakes` Table. For each sample determine the maximum value of `mag` in the sample. All 5,000 maximum values from the samples should be stored in an array named `maximums`.


```python
maximums = ...
for i in np.arange(5000): 
    ...
```

```python
grader.check("q19")
```

Run the cell below to create a histogram of the 5,000 maximums you simulated to view their distribution. You'll need this to help answer question 4.5.

```python
# Histogram of your maximums
Table().with_column('Largest magnitude in sample', maximums).hist('Largest magnitude in sample') 
```

**Question 20.** Now find the magnitude of the **actual** strongest earthquake in 2019 using the `earthquake` Table, not the maximum of just a sample.


```python
strongest_earthquake_magnitude = ...
strongest_earthquake_magnitude
```

```python
grader.check("q20")
```

<!-- BEGIN QUESTION -->

**Question 21.**  Explain whether you believe you can accurately use a sample size of 200 to determine the maximum. Use the histogram and population maximum of `mag` determined in **Question 19.** and **Question 20.** to help answer this question. What is one problem with using the maximum function as your estimator that you don't have to worry about when using an average?

**Hint:** Use the histogram to estimate how many of the samples got the maximum correct?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

## 5. Assessing Jade's Models
#### Games with Jade

Our friend Jade comes over and asks us to play a game with her. The game works like this: 

> We will draw randomly with replacement from a simplified 13 card deck with 4 face cards (A, J, Q, K), and 9 numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10). If we draw cards with replacement 13 times, and if the number of face cards is greater than or equal to 4, we lose.
> 
> Otherwise, Jade wins.

We play the game once and we lose, observing 8 total face cards. We are angry and accuse Jade of cheating! Jade is adamant, however, that the deck is fair.

#### Jade's model of the game
Jade claims that there is an equal chance of getting any of the cards (A, 2, 3, 4, 5, 6, 7, 8, 9, 10, J, Q, K), but we do not believe her. 

#### Our alternative model of the game
We believe that the deck is clearly rigged, with face cards (A, J, Q, K) being more likely than the numbered cards (2, 3, 4, 5, 6, 7, 8, 9, 10).

**Question 22.** Assign `deck_model_probabilities` to a two-item array containing the chance of drawing a face card as the first element, and the chance of drawing a numbered card as the second element under the assumptions of Jade's model. Since we're working with probabilities, make sure your values are between 0 and 1. Probabilities should be exact representations of the values (1/3 not 0.333).


```python
deck_model_probabilities = ...
deck_model_probabilities
```

```python
grader.check("q22")
```

**Question 23.** We believe Jade's model (every card is equally likely to be drawn) is incorrect. In particular, we believe the deck is rigged such that there is a larger chance of getting a face card.  Which of the following statistics can we calculate in our simulation to test for differences between the model and our alternative? Assign `statistic_choice` to the correct answer. 

**Hint:** Look back at the data we observed from the game we've already played. 

1. The actual number of face cards we get in 13 draws

1. The distance (absolute value) between the actual number of face cards in 13 draws and the expected number of face cards in 13 draws (4)

1. The expected number of face cards in 13 draws (4)


```python
statistic_choice = ...
statistic_choice
```

```python
grader.check("q23")
```

<!-- BEGIN QUESTION -->

**Question 24.** Define the function `deck_simulation_and_statistic`, which, given an integer sample size and an array of model proportions (like the one you created in **Question 22**), returns **the number of face cards** in one simulation in of drawing cards under the model specified in `model_proportions`. The included final line of code in the cell below will call your function to simulate drawing 13 cards from the deck using the probabilities assigned to the array `deck_model_probabilities` and return how many face cards those 13 drawn cards contained.

**Hint:** Think about how you can use the function `sample_proportions` contained in the `datascience` library. 


```python
def deck_simulation_and_statistic(sample_size, model_proportions):
    ...

deck_simulation_and_statistic(13, deck_model_probabilities)
```

```python
grader.check("q24")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 25.** Use your function from **Question 24.** to run 5,000 simulations in which 13 cards are drawn under the proportions that you specified in **Question 22.** Store of all of your statistics in an array named `deck_statistics`. 


```python
repetitions = 5000 
...

deck_statistics
```

```python
grader.check("q25")
```

<!-- END QUESTION -->

Let’s take a look at the distribution of simulated statistics. The cell below will create a histogram of your results. How many face cards would seem to be the most common result when dealt 13 cards when assuming our proportions from **Question 22**? You'll want to keep that in mind when answering the next question.

```python
# Draw a distribution of statistics 
Table().with_column('Deck Statistics', deck_statistics).hist(bins=np.arange(-0.5,13.5,1))
```

<!-- BEGIN QUESTION -->

**Question 26.** Given the observed value of being dealt 8 face cards out of 13 total cards, do you believe that Jade's model is reasonable? Explain your answer using the distribution drawn in the previous problem. In particular, consider how likely such a result would be using the information in the histogram.


_Type your answer here, replacing this text._

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: None)
```
>>> ten_sixes in [1, 2, 3]
True
```


Tests for q10 (points: None)
```
>>> round(m_chance, 5)
0.09295
```


Tests for q12 (points: None)
```
>>> full_data.num_rows == 562
True
```
```
>>> list(full_data.labels)[0] == 'Player'
True
```
```
>>> full_data.select(sorted(full_data.labels)).sort(4).take(range(3))
2P   | 3P   | PTS  | Player         | Salary
0.6  | 0    | 1.7  | Tyler Cook     | 50000
0    | 0    | 0    | William Howard | 50000
2    | 0    | 6    | Eric Mika      | 50752
```


Tests for q15 (points: None)
```
>>> distribution_1 in ['empirical', 'probability']
True
```
```
>>> distribution_2 in ['empirical', 'probability']
True
```
```
>>> type(player_count_1) == int and type(player_count_2) == int
True
```
```
>>> type(area_total_1) == int and type(area_total_2) == int
True
```


Tests for q18 (points: None)
```
>>> representative_sample.num_rows == 200
True
```
```
>>> all(np.in1d(representative_sample.column('mag'), earthquakes.column('mag')))
True
```
```
>>> representative_mean < max(representative_sample.column('mag')) and representative_mean > min(representative_sample.column('mag'))
True
```


Tests for q19 (points: None)
```
>>> len(maximums) != 0
True
```
```
>>> len(maximums) == 5000
True
```
```
>>> max(maximums) <= max(earthquakes.column('mag'))
True
```


Tests for q2 (points: None)
```
>>> five_or_less in [1, 2, 3]
True
```


Tests for q20 (points: None)
```
>>> isinstance(strongest_earthquake_magnitude, float)
True
```


Tests for q22 (points: None)
```
>>> len(deck_model_probabilities) == 2
True
```
```
>>> sum(deck_model_probabilities) == 1
True
```


Tests for q23 (points: None)
```
>>> statistic_choice in [1, 2, 3]
True
```


Tests for q24 (points: None)
```
>>> all([0 <= deck_simulation_and_statistic(13, deck_model_probabilities) <= 13 for k in range(50)])
True
```


Tests for q25 (points: None)
```
>>> len(deck_statistics) == repetitions
True
```
```
>>> all([0 <= k <= 13 for k in deck_statistics])
True
```


Tests for q3 (points: None)
```
>>> lottery in [1, 2, 3]
True
```


Tests for q4 (points: None)
```
>>> list_chances in [1, 2, 3]
True
```


Tests for q5 (points: None)
```
>>> np.isclose(ma4110_chance, (1 / 62) ** 6)
True
```


Tests for q6 (points: None)
```
>>> simulate_key_strike() is not None
True
```
```
>>> all([simulate_key_strike() in list(string.ascii_lowercase + string.ascii_uppercase + string.digits) for i in range(100)])
True
```
```
>>> np.random.seed(22)
>>> 62 >= len(np.unique([simulate_key_strike() for i in range(500)])) >= 45
True
```


Tests for q7 (points: None)
```
>>> len(simulate_several_key_strikes(15))
15
```
```
>>> isinstance(simulate_several_key_strikes(15), str)
True
```
```
>>> np.random.seed(22)
>>> 62 >= len(np.unique(list(simulate_several_key_strikes(500)))) >= 45
True
```


Tests for q8 (points: None)
```
>>> ma4110_proportion == 0.0
True
```
